```python
# Initialize Otter
import otter
grader = otter.Notebook("lab09.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

<div style="text-align: center;">
    <h1>Lab 09: Sample Means</h1>
    <em>View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.</em>
</div>

In this lab, you will learn about [the variance of sample means](https://inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html) as well as [the normal distribution](https://www.inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html).

Run the following cell to to import the modules and settings for this lab.

```python
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore')
```

## Normal Distributions

When we visualize the distribution of a sample, we are often interested in the mean and the standard deviation of the sample (for the rest of this lab, we will abbreviate “standard deviation” as “SD”). These two summary statistics can give us a bird’s eye view of the distribution - by letting us know where the distribution sits on the number line and how spread out it is, respectively. 

We want to check if the data is linearly related, so we should look at the data.

### Task 01 📍🔎

The next cell loads the table `births` from lecture, which is a large random sample of US births and includes information about mother-child pairs. 

Plot the distribution of mother's ages from the table. 

Don't change the last line, which will plot the mean of the sample on the distribution itself.

_Check your histogram with someone before moving on._


```python
births = Table.read_table('baby.csv')
...

# Do not change this line
plots.scatter(np.mean(births.column("Maternal Age")), 0, color='red', zorder=2, s=50);
plots.ylim(-0.005, 0.07);
```

From the plot above, we can see that the mean is the center of gravity or balance point of the distribution. If you cut the distribution out of cardboard, and then placed your finger at the mean, the distribution would perfectly balance on your finger. Since the distribution above is right skewed (which means it has a long right tail), we know that the mean of the distribution is larger than the median, which is the “halfway” point of the data. Conversely, if the distribution had been left skewed, we know the mean would be smaller than the median.

### Task 02 📍🔎

<!-- BEGIN QUESTION -->

1. Run the following cell to compare the mean (red) and median (green) of the distribution of mothers ages.
2. Why do you think the mean is positioned to the right of the median.

_Check your response with someone before moving on._

```python
births.hist("Maternal Age")
plots.scatter(np.mean(births.column("Maternal Age")), 0, color='red', zorder=2, s=50, label="Mean")
plots.scatter(np.median(births.column("Maternal Age")), 0, color='blue', zorder=2, s=50, label="Median")
plots.legend()
plots.ylim(-0.005, 0.07);
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

We are also interested in the standard deviation of mother’s ages. The SD gives us a sense of how variable mothers' ages are around the average mothers' age. If the SD is large, then the mothers' heights should spread over a large range from the mean. If the SD is small, then the mothers' heights should be tightly clustered around the average mother height. 

**The SD of an array is defined as the root mean square of deviations (differences) from average**.

Fun fact! σ (Greek letter sigma) is used to represent the SD and  μ (Greek letter mu) is used for the mean.

### Task 03 📍

1. Define `age_mean` to be the mean of the maternal ages.
2. Define `age_sd` to be the standard deviation of the maternal ages.
3. Run the cell below to see the width of one SD (blue) from the sample mean (red) plotted on the histogram of maternal ages.

```python
age_mean = ...
age_sd = ...
births.hist("Maternal Age")

plots.scatter(age_mean, 0, color='red', zorder=2, s=50);
plots.scatter(age_mean+age_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.scatter(age_mean-age_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.ylim(-0.005, 0.07);
```

```python
grader.check("task_03")
```

In this histogram, the standard deviation is not easy to identify just by looking at the graph.

However, the distributions of some variables allow us to easily spot the standard deviation on the plot. For example, if a sample follows a *normal distribution*, the standard deviation is easily spotted at the point of inflection (the point where the curve begins to change the direction of its curvature) of the distribution. 

### Task 04 📍

1. Define `height_mean` to be the mean height value.
2. Define `height_sd` to be the standard deviation of the height values.
3. Examine the distribution of maternal heights, which is roughly normally distributed. The code will plot the standard deviation on the histogram, as before - notice where one standard deviation (blue) away from the mean (red) falls on the plot.


```python
height_mean = ...
height_sd = ...
births.hist("Maternal Height", bins=np.arange(55,75,1))

plots.scatter((height_mean), 0, color='red', zorder=2, s=50);
plots.scatter(height_mean+height_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.scatter(height_mean-height_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.ylim(-0.01, 0.16);
```

```python
grader.check("task_04")
```

We don't always know how a variable will be distributed, and making assumptions about whether or not a variable will follow a normal distribution is dangerous. However, the Central Limit Theorem defines one distribution that always follows a normal distribution. The distribution of the *sums* and *means* of many large random samples drawn with replacement from a single distribution (regardless of the distribution's original shape) will be normally distributed. Remember that the Central Limit Theorem refers to the distribution of a *statistic* calculated from a distribution, not the distribution of the original sample or population.

The next section will explore distributions of sample means, and you will see how the standard deviation of these distributions depends on sample sizes.

## Variability of the Sample Mean

By the [Central Limit Theorem](https://www.inferentialthinking.com/chapters/14/4/Central_Limit_Theorem.html), the probability distribution of the mean of a large random sample is roughly normal. The bell curve is centered at the population mean. Some of the sample means are higher and some are lower, but the deviations from the population mean are roughly symmetric on either side, as we have seen repeatedly. Formally, probability theory shows that the sample mean is an **unbiased estimate** of the population mean.

In our simulations, we also noticed that the means of larger samples tend to be more tightly clustered around the population mean than means of smaller samples. In this section, we will quantify the [variability of the sample mean](https://www.inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html) and develop a relation between the variability and the sample size.

Let's take a look at the salaries of employees of the City of San Francisco in 2019. The mean salary reported by the city government was about $80,608.71.

*Note: If you get stuck on any part of this lab, please refer to [chapter 14 of the textbook](https://www.inferentialthinking.com/chapters/14/Why_the_Mean_Matters.html).*

```python
salaries = Table.read_table('san_francisco_2019.csv').select("Salary").relabeled("Salary", "salary")
salaries
```

```python
salary_mean = np.mean(salaries.column('salary'))
print(f'Mean salary of San Francisco city employees in 2019: ${round(salary_mean, 2)}')
```

```python
salaries.hist('salary', bins=np.arange(0, 250000+10000*2, 10000))
plots.scatter(salary_mean, 0, color='red', zorder=2, s=50)
plots.title('2019 salaries of city of SF employees')
plots.ylim(-0.0000005, 0.000015);
```

Clearly, the population does not follow a normal distribution. Keep that in mind as we progress through these exercises.

Let's take random samples *with replacement* and look at the probability distribution of the sample mean. As usual, we will use simulation to get an empirical approximation to this distribution.

### Task 05 📍

Define a function `one_sample_mean`. Its arguments should be `table` (the name of a table), `label` (the label of the column containing the variable), and `sample size`(the number of employees in the sample). It should sample with replacement from the table and
return the mean of the `label` column of the sample.


```python
def one_sample_mean(table, label, sample_size):
    new_sample = ...
    new_sample_mean = ...
    ...
```

```python
grader.check("task_05")
```

### Task 06 📍🔎

Use `one_sample_mean` to define a function `simulate_sample_mean`. The arguments are the name of the table, the label of the column containing the variable, the sample size, and the number of simulations.

The function should sample with replacement from the table and calculate the mean of each sample. It should save the sample means in an array called `means`. The remaining code in the function displays an empirical histogram of the sample means.


```python
"""Empirical distribution of random sample means"""

def simulate_sample_mean(table, label, sample_size, repetitions):
    
    means = make_array()

    for i in np.arange(repetitions):
        new_sample_mean = ...
        means = ...

    sample_means = Table().with_column('Sample Means', means)
    
    # Display empirical histogram and print all relevant quantities – don't change this!
    sample_means.hist(bins=20)
    plots.xlabel('Sample Means')
    plots.title('Sample Size ' + str(sample_size))
    print("Sample size: ", sample_size)
    print("Population mean:", np.mean(table.column(label)))
    print("Average of sample means: ", np.mean(means))
    print("Population SD:", np.std(table.column(label)))
    print("SD of sample means:", np.std(means))
    return np.std(means)
```

---

**Verify with someone that you've implemented the function above correctly. If you haven't implemented it correctly, the rest of the lab won't work properly, so this step is crucial.**

 In the following cell, we will create a sample of size 100 from `salaries` and graph it using our new `simulate_sample_mean` function.
 
 *Hint: You should see a distribution similar to something we've been talking about. If not, check your function*

---

```python
simulate_sample_mean(salaries, 'salary', 100, 10000) 
plots.xlim(50000, 100000);
```

### Task 07 📍🔎

Simulate two sample means, one for a sample of 400 salaries and one for a sample of 625 salaries. In each case, perform 10,000 repetitions. 

* _Don't worry about the `plots.xlim` line – it just makes sure that all of the plots have the same x-axis._
* _Compare your histograms with someone else before moving on._


```python
simulate_sample_mean(..., ..., ..., ...)
plots.xlim(50000, 100000);
plots.show();
print('\n')
simulate_sample_mean(..., ..., ..., ...)
plots.xlim(50000, 100000);
plots.show();

```

### Task 08 📍

Assign `task_08_statements` to an array of numbers corresponding to true statement(s) about the plots from in Task 07.

1. One of the conditions for CLT is that we have to draw a small random sample with replacement from the population.
1. One of the conditions for CLT is that we have to draw a large random sample with replacement from the population.
1. One of the conditions for CLT is that the population must be normally distributed.
1. Both plots in Task 07 are roughly centered around the population mean.
1. Both plots in Task 07 are roughly centered around the mean of a particular sample.
1. The distribution of sample means for sample size 625 has less variability than the distribution of sample means for sample size 400.
1. The distribution of sample means for sample size 625 has more variability than the distribution of sample means for sample size 400.
1. We see the Central Limit Theorem (CLT) in action because the distributions of the sample means are bell-shaped.
1. We see the Law of Averages in action because the distributions of the sample means look like the distribution of the population.


```python
task_08_statements = ...
```

```python
grader.check("task_08")
```

### Task 09 📍🔎

Below, we'll look at what happens when we take an increasing number of resamples of a fixed sample size. Notice what number in the code changes, and what stays the same. How does the distribution of the resampled means change?

```python
simulate_sample_mean(salaries, 'salary', 100, 500)
plots.xlim(50000, 100000);
```

```python
simulate_sample_mean(salaries, 'salary', 100, 1000)
plots.xlim(50000, 100000);
```

```python
simulate_sample_mean(salaries, 'salary', 100, 5000)
plots.xlim(50000, 100000);
```

```python
simulate_sample_mean(salaries, 'salary', 100, 10000)
plots.xlim(50000, 100000);
```

What did you notice about the distributions of sample means in the four histograms above? 

_Check your response with someone before moving on._

_Type your answer here, replacing this text._

### Task 10 📍

Assign the variable `SD_of_sample_means` to the integer corresponding to your answer to the following question:

When I increase the number of resamples that I take, for a fixed sample size, the SD of my sample means will...

1. Increase
2. Decrease
3. Stay about the same
4. Vary widly



```python
SD_of_sample_means = ...
```

```python
grader.check("task_10")
```

### Task 11 📍

Let's think about how the relationships between population SD, sample SD, and SD of sample means change with varying sample size. Which of the following is true? Assign the variable `pop_vs_sample` to an array of integer(s) that correspond to true statement(s).

1. Sample SD gets smaller with increasing sample size.
2. Sample SD gets larger with increasing sample size.
3. Sample SD becomes more consistent with population SD with increasing sample size.
4. SD of sample means gets smaller with increasing sample size.
5. SD of sample means gets larger with increasing sample size.
6. SD of sample means stays the same with increasing sample size.


```python
pop_vs_sample = ...
```

```python
grader.check("task_11")
```

Run the following three cells multiple times and examine how the sample SD and the SD of sample means change with sample size.

The first histogram is of the sample; the second histogram is the distribution of sample means with that particular sample size. Adjust the bins as necessary.

```python
sample_10 = salaries.sample(10)
sample_10.hist("salary")
plots.title('Distribution of salary for sample size 10')
print("Sample SD: ", np.std(sample_10.column("salary")))
simulate_sample_mean(salaries, 'salary', 10, 1000)
plots.xlim(5,120000);
plots.ylim(0, .0001);
plots.title('Distribution of sample means for sample size 10');
```

```python
sample_200 = salaries.sample(200)
sample_200.hist("salary")
plots.title('Distribution of salary for sample size 200')
print("Sample SD: ", np.std(sample_200.column("salary")))
simulate_sample_mean(salaries, 'salary', 200, 1000)
plots.xlim(5,100000)
plots.ylim(0, .00015);
plots.title('Distribution of sample means for sample size 200');
```

```python
sample_1000 = salaries.sample(1000)
sample_1000.hist("salary")
plots.title('Distribution of salary for sample size 1000')
print("Sample SD: ", np.std(sample_1000.column("salary")))
simulate_sample_mean(salaries, 'salary', 1000, 1000)
plots.xlim(5,100000)
plots.ylim(0, .00025);
plots.title('Distribution of sample means for sample size 1000');
```

You should notice that the distribution of means gets narrower and spikier, and that the distribution of the sample increasingly looks like the distribution of the population as we get to larger sample sizes. 

Let's illustrate these trends. Below, you will see how the sample SD changes with respect to sample size (N). The blue line is the population SD.

```python
# Don't change this cell, just run it!
pop_sd = np.std(salaries.column('salary'))
sample_sds = make_array()
sample_sizes = make_array()
for i in np.arange(10, 500, 10):
    sample_sds = np.append(sample_sds, [np.std(salaries.sample(i).column("salary")) for d in np.arange(100)])
    sample_sizes = np.append(sample_sizes, np.ones(100) * i)
Table().with_columns("Sample SD", sample_sds, "N", sample_sizes).scatter("N", "Sample SD")
matplotlib.pyplot.axhline(y=pop_sd, color='blue', linestyle='-');
```

The next cell shows how the SD of the sample means changes relative to the sample size (N).

```python
# Don't change this cell, just run it!
def sample_means(sample_size):
    means = make_array()
    for i in np.arange(1000):
        sample = salaries.sample(sample_size).column('salary')
        means = np.append(means, np.mean(sample))
    return np.std(means)

sample_mean_SDs = make_array()
for i in np.arange(50, 1000, 100):
    sample_mean_SDs = np.append(sample_mean_SDs, sample_means(i))
Table().with_columns("SD of sample means", sample_mean_SDs, "Sample Size", np.arange(50, 1000, 100))\
.plot("Sample Size", "SD of sample means")
```

From these two plots, we can see that the SD of our *sample* approaches the SD of our population as our sample size increases, but the SD of our *sample means* (in other words, the variability of the sample mean) decreases as our sample size increases.

### Task 12 📍

Is there a relationship between the sample size and the standard deviation of the sample mean? Assign `task_12` to the number corresponding to the statement that answers this question.

1. The SD of the sample means is inversely proportional to the square root of sample size.
2. The SD of the sample means is directly proportional to the square root of sample size.


```python
task_12 = ...
```

```python
grader.check("task_12")
```

Throughout this lab, we have been taking many random samples from a population. However, all of these principles hold for bootstrapped resamples from a single sample. If your original sample is relatively large, all of your re-samples will also be relatively large, and so the SD of resampled means will be relatively small. 

In order to change the variability of your sample mean, you’d have to change the size of the original sample from which you are taking bootstrapped resamples.

## Submit your Lab to Canvas

Once you have finished working on the lab questions, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the requirements for a Complete score for this lab assignments.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
4. Select the menu items "File", "Save and Export Notebook As...", and "HTML (.html)" in the notebook's Toolbar to download an HTML version of this notebook file.
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded HTML file.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_03 (points: None)
```
>>> 
>>> np.isclose(round(age_mean, 3), 27.228)
True
```
failure_message: ❌ Your mean value doesn't seem correct. Use np.mean on the array of maternal ages.
success_message: ✅ Your mean value seems correct.
```
>>> 
>>> np.isclose(round(age_sd, 3), 5.815)
True
```
failure_message: ❌ Your standard deviation value doesn't seem correct. Use np.std on the array of maternal ages.
success_message: ✅ Your standard deviation value seems correct.


Tests for task_04 (points: None)
```
>>> 
>>> np.isclose(round(height_mean, 3), 64.049)
True
```
failure_message: ❌ Your mean value doesn't seem correct. Use np.mean on the array of heights.
success_message: ✅ Your mean value seems correct.
```
>>> 
>>> np.isclose(round(height_sd, 3), 2.525)
True
```
failure_message: ❌ Your standard deviation value doesn't seem correct. Use np.std on the array of heights.
success_message: ✅ Your standard deviation value seems correct.


Tests for task_05 (points: None)
```
>>> 
>>> np.random.seed(8)
>>> round(one_sample_mean(salaries, 'salary', 100), 11) == 82072.729999999996
True
```
failure_message: ❌ one_sample_mean didn't work for the inputs that we checked. Make sure you are using the sample method, specifying the sample size, and using with_replacement.
success_message: ✅ one_sample_mean worked for the inputs that we checked.


Tests for task_08 (points: None)
```
>>> 
>>> set(task_08_statements) == set([2, 4, 6, 8])
True
```
failure_message: ❌ Some of your selections are not correct. Try again.
success_message: ✅ Great work!.


Tests for task_10 (points: None)
```
>>> 
>>> SD_of_sample_means == 3
True
```
failure_message: ❌ That is not correct.
success_message: ✅ That is correct.


Tests for task_11 (points: None)
```
>>> 
>>> set(pop_vs_sample) == set([3,4])
True
```
failure_message: ❌ That is not correct.
success_message: ✅ That is correct.


Tests for task_12 (points: None)
```
>>> 
>>> task_12 == 1
True
```
failure_message: ❌ That is not correct.
success_message: ✅ That is correct.
