```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

<div style="text-align: center;">
    <h1>Homework 02 - Arrays and Tables</h1>
    <em>View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.</em>
</div>

```python
# Import Modules
import numpy as np
from datascience import *
```

**Recommended Reading**: 
* [Data Types](https://inferentialthinking.com/chapters/04/Data_Types.html) 
* [Sequences](https://inferentialthinking.com/chapters/05/Sequences.html)
* [Tables](https://inferentialthinking.com/chapters/06/Tables.html)

Please complete this notebook by filling in the cells provided.
 
Throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. 

## Creating Arrays


### Note About Arrays for Those with Python Experience

Python lists are different/behave differently than NumPy arrays. In this course, we use a version of NumPy arrays, so please make your arrays using tools in the `datascience` module such as `make_array` and don't use a python list.

### Task 01 📍

Make an array called `weird_numbers` containing the following numbers (in the given order):

1. 3,000
2. the square root of 5
3. -2
4. $\pi$ to the power of the square root of 2

Remember that `math` module contains the square root function and the constant value $\pi$.

_Points:_ 2

```python
# You will probably want to import something to start!
...
weird_numbers = ...
weird_numbers
```

```python
grader.check("task_01")
```

### Task 02 📍

Make an array called `book_title_words` containing the following three strings: "Eats", "Shoots", and "and Leaves".


_Points:_ 2

```python
book_title_words = ...
book_title_words
```

```python
grader.check("task_02")
```

Strings have a method called `join`.  `join` takes one argument, an array of strings.  It returns a single string.  Specifically, the value of `a_string.join(an_array)` is a single string that's the [concatenation](https://en.wikipedia.org/wiki/Concatenation) ("putting together") of all the strings in `an_array`, **except** `a_string` is inserted in between each string.

### Task 03 📍

Use the array `book_title_words` and the method `join` to make two strings:

1. "Eats, Shoots, and Leaves" (call this one `with_commas`)
2. "Eats Shoots and Leaves" (call this one `without_commas`)

*Hint:* If you're not sure what `join` does, first try just calling, for example, `"foo".join(book_title_words)` .


_Points:_ 2

```python
with_commas = ...
without_commas = ...

# These lines are provided just to print out your answers.
print('with_commas:', with_commas)
print('without_commas:', without_commas)
```

```python
grader.check("task_03")
```

## Indexing Arrays


These exercises give you practice accessing individual elements of arrays.  In Python (and in many programming languages), elements are accessed by *index*, so the first element is the element at index 0, the second element is at index 1, etc.

### Note About Indexing for Those with Python Experience

Please don't use bracket notation when indexing (i.e. `arr[0]`), as this can yield different data type outputs than what we will be expecting. This can cause you to fail an autograder test.  Instead, use `arr.item(0)`.

### Task 04 📍

The cell below creates an array of some numbers.  Set `third_element` to the third element of `some_numbers`.


_Points:_ 2

```python
some_numbers = make_array(-1, -3, -6, -10, -15)

third_element = ...
third_element
```

```python
grader.check("task_04")
```

The test above checks that your answer is in the correct format. **This test does not check that you answered correctly**, only that you assigned a number successfully. The same is true for most of the tests in this homework, and every other homework for this class.

### Task 05 📍

The next cell creates a table that displays some information about the elements of `some_numbers` and their order.  Run the cell to see the partially-completed table, then fill in the missing information (the cells that say "Ellipsis") by assigning `blank_a`, `blank_b`, `blank_c`, and `blank_d` to the correct elements in the table.


_Points:_ 2

```python
blank_a = ...
blank_b = ...
blank_c = ...
blank_d = ...
elements_of_some_numbers = Table().with_columns(
    "English name for position", make_array("first", "second", blank_a, blank_b, "fifth"),
    "Index",                     make_array(blank_c, 1, 2, blank_d, 4),
    "Element",                   some_numbers)
elements_of_some_numbers
```

```python
grader.check("task_05")
```

### Task 06 📍

You'll sometimes want to find the *last* element of an array.  Suppose an array has 145 elements.  What is the index of its last element?


_Points:_ 2

```python
index_of_last_element = ...
```

```python
grader.check("task_06")
```

More often, you don't know the number of elements in an array, its *length*.  (For example, it might be a large dataset you found on the Internet.)  The function `len` takes a single argument, an array, and returns the `len`gth of that array (an integer).

### Task 07 📍

The cell below loads an array called `president_birth_years`.  Calling `.column(...)` on a table returns an array of the column specified, in this case the `Birth Year` column of the `president_births` table. The last element in that array is the most recent birth year of any deceased president. Assign that year to `most_recent_birth_year`.


_Points:_ 2

```python
president_birth_years = Table.read_table("president_births.csv").column('Birth Year')

most_recent_birth_year = ...
most_recent_birth_year
```

```python
grader.check("task_07")
```

### Task 08 📍

Finally, assign `sum_of_birth_years` to the sum of the third, tenth, and last birth year in `president_birth_years`. You might consider breaking up your solution into parts.

_Points:_ 2

```python
sum_of_birth_years = ...
```

```python
grader.check("task_08")
```

## Basic Array Arithmetic


### Task 09 📍

Multiply the numbers 42, 4224, 42422424, and -250 by 157. Assign each variable below such that `first_product` is assigned to the result of $42 * 157$, `second_product` is assigned to the result of $4224 * 157$, and so on. 

For this task, **don't** use arrays. You'll use them later.


_Points:_ 2

```python
first_product = ...
second_product = ...
third_product = ...
fourth_product = ...
print(first_product, second_product, third_product, fourth_product)
```

```python
grader.check("task_09")
```

### Task 10 📍

Now, do the same calculation, but using an array called `numbers` and only a single multiplication (`*`) operator.  Store the 4 results in an array named `products`.


_Points:_ 2

```python
numbers = ...
products = numbers * 157
products
```

```python
grader.check("task_10")
```

### Task 11 📍

Oops, we made a typo!  Instead of 157, we wanted to multiply each number by 1577.  Compute the correct products in the cell below using array arithmetic.  Notice that your job is really easy if you previously defined an array containing the 4 numbers.


_Points:_ 2

```python
correct_products = ...
correct_products
```

```python
grader.check("task_11")
```

### Task 12 📍

We've loaded an array of temperatures in the next cell.  Each number is the highest temperature observed on a particular day at a climate observation station, mostly from the US.  Since they're from the US government agency [NOAA](https://noaa.gov), all the temperatures are in Fahrenheit.

Convert all temperature values from Fahrenheit to Celsius by first subtracting 32 from them, then multiplying the results by $\frac{5}{9}$. Make sure to **ROUND** the final result after converting to Celsius to the nearest integer using the `np.round` function. Remember to check the documentation of a function if you are not sure how it works.

_Points:_ 2

```python
max_temperatures = Table.read_table("temperatures.csv").column("Daily Max Temperature")

celsius_max_temperatures = ...
celsius_max_temperatures
```

```python
grader.check("task_12")
```

### Task 13 📍

The cell below loads all the *lowest* temperatures from each day (in Fahrenheit).  Compute the size of the daily temperature range for each day.  That is, compute the difference between each daily maximum temperature and the corresponding daily minimum temperature.  **Pay attention to the units, give your answer in Celsius!** Make sure **NOT** to round your answer for this question!


_Points:_ 2

```python
min_temperatures = Table.read_table("temperatures.csv").column("Daily Min Temperature")

celsius_temperature_ranges = ...
celsius_temperature_ranges
```

```python
grader.check("task_13")
```

## World Population


The cell below loads a table of estimates of the world population for from 1950 until 2015. The estimates come from the [US Census Bureau website](https://www.census.gov/en.html).

```python
world = Table.read_table("world_population.csv").select('Year', 'Population')
world
```

The name `population` is assigned to an array of population estimates.

```python
population = world.column(1)
population
```

In this question, you will apply some built-in Numpy functions to this array. Numpy is a module that is often used in Data Science!

<img src="array_diff.png" style="width: 600px;"/>

The difference function `np.diff` subtracts each element in an array from the element after it within the array. As a result, the length of the array `np.diff` returns will always be one less than the length of the input array.

<img src="array_cumsum.png" style="width: 700px;"/>

The cumulative sum function `np.cumsum` outputs an array of partial sums. For example, the third element in the output array corresponds to the sum of the first, second, and third elements.

### Task 14 📍

Very often in data science, we are interested understanding how values change with time. Use `np.diff` and `np.max` (or just `max`) to calculate the largest annual change in population between any two consecutive years.


_Points:_ 2

```python
largest_population_change = ...
largest_population_change
```

```python
grader.check("task_14")
```

### Task 15 📍

What do the values in the resulting array represent (choose one)?

```python
np.cumsum(np.diff(population))
```

1) The total population change between consecutive years, starting at 1951.

2) The total population change between 1950 and each later year, starting at 1951.

3) The total population change between 1950 and each later year, starting inclusively at 1950.

Assign `cumulative_sum_answer` to 1, 2, or 3


_Points:_ 2

```python
cumulative_sum_answer = ...
```

```python
grader.check("task_15")
```

## Old Faithful


Old Faithful is a geyser in Yellowstone that erupts every 44 to 125 minutes (according to [Wikipedia](https://en.wikipedia.org/wiki/Old_Faithful)). People are [often told that the geyser erupts every hour](http://yellowstone.net/geysers/old-faithful/), but in fact the waiting time between eruptions is more variable. Let's take a look.

### Task 16 📍

The first line below assigns `waiting_times` to an array of 272 consecutive waiting times between eruptions, taken from a classic 1938 dataset. Assign the names `shortest`, `longest`, and `average` so that the `print` statement is correct.


_Points:_ 2

```python
waiting_times = Table.read_table('old_faithful.csv').column('waiting')

shortest = ...
longest = ...
average = ...

print("Old Faithful erupts every", shortest, "to", longest, "minutes and every", average, "minutes on average.")
```

```python
grader.check("task_16")
```

### Task 17 📍

Assign `biggest_decrease` to the biggest decrease in waiting time between two consecutive eruptions. For example, the third eruption occurred after 74 minutes and the fourth after 62 minutes, so the decrease in waiting time was 74 - 62 = 12 minutes. 

*Hint 1*: You'll need an array arithmetic function [mentioned in the textbook](https://inferentialthinking.com/chapters/05/1/Arrays.html#functions-on-arrays). You have also seen this function earlier in the homework!

*Hint 2*: We want to return the absolute value of the biggest decrease.


_Points:_ 2

```python
biggest_decrease = ...
biggest_decrease
```

```python
grader.check("task_17")
```

### Task 18 📍

If you expected Old Faithful to erupt every hour, you would expect to wait a total of `60 * k` minutes to see `k` eruptions. Set `difference_from_expected` to an array with 272 elements, where the element at index `i` is the absolute difference between the expected and actual total amount of waiting time to see the first `i+1` eruptions. 

*Hint*: You'll need to compare a cumulative sum to a range of values. We recommend that you use `np.arange`, and you can read about it in[section 5.2 of the textbook](https://www.inferentialthinking.com/chapters/05/2/Ranges.html).

For example, since the first three waiting times are 79, 54, and 74, the total waiting time for 3 eruptions is 79 + 54 + 74 = 207. The expected waiting time for 3 eruptions is 60 * 3 = 180. Therefore, `difference_from_expected.item(2)` should be $|207 - 180| = 27$.


_Points:_ 1

```python
difference_from_expected = ...
difference_from_expected
```

```python
grader.check("task_18")
```

### Task 19 📍

Let’s imagine your guess for the next wait time was always just the length of the previous waiting time. If you always guessed the previous waiting time, how big would your error in guessing the waiting times be, on average?

For example, since the first three waiting times are 79, 54, and 74, the average difference between your guess and the actual time for just the second and third eruption would be $\frac{|79-54|+ |54-74|}{2} = 22.5$.


_Points:_ 2

```python
average_error = ...
average_error
```

```python
grader.check("task_19")
```

## Tables


### Task 20 📍

Suppose you have 4 apples, 3 oranges, and 3 pineapples.  (Perhaps you're using Python to solve a high school Algebra problem.)  Create a table that contains this information.  It should have two columns: `fruit name` and `count`.  Assign the new table to the variable `fruits`.

**Note:** Use lower-case and singular words for the name of each fruit, like `"apple"`.


_Points:_ 2

```python
# Our solution uses 1 statement split over 3 lines.
fruits = ...
        ...
        ...
fruits
```

```python
grader.check("task_20")
```

### Task 21 📍

The file `inventory.csv` contains information about the inventory at a fruit stand.  Each row represents the contents of one box of fruit. Load it as a table named `inventory` using the `Table.read_table()` function. `Table.read_table(...)` takes one argument (data file name in string format) and returns a table.


_Points:_ 1

```python
inventory = ...
inventory
```

```python
grader.check("task_21")
```

### Task 22 📍

Does each box at the fruit stand contain a different fruit? Set `all_different` to `True` if each box contains a different fruit or to `False` if multiple boxes contain the same fruit.

*Hint:* You don't have to write code to calculate the True/False value for `all_different`. Just look at the `inventory` table and assign `all_different` to either `True` or `False` according to what you can see from the table in answering the question.


_Points:_ 2

```python
all_different = ...
all_different
```

```python
grader.check("task_22")
```

### Task 23 📍

The file `sales.csv` contains the number of fruit sold from each box last Saturday.  It has an extra column called "price per fruit (\$)" that's the price *per item of fruit* for fruit in that box.  The rows are in the same order as the `inventory` table.  Load these data into a table called `sales`.


_Points:_ 1

```python
sales = ...
sales
```

```python
grader.check("task_23")
```

### Task 24 📍

How many fruits did the store sell in total on that day?


_Points:_ 2

```python
total_fruits_sold = ...
total_fruits_sold
```

```python
grader.check("task_24")
```

### Task 25 📍

What was the store's total revenue (the total price of all fruits sold) on that day?

*Hint:* If you're stuck, think first about how you would compute the total revenue from just the grape sales.


_Points:_ 2

```python
total_revenue = ...
total_revenue
```

```python
grader.check("task_25")
```

### Task 26 📍

Make a new table called `remaining_inventory`.  It should have the same rows and columns as `inventory`, except that the amount of fruit sold from each box should be subtracted from that box's count, so that the "count" is the amount of fruit remaining after Saturday.


_Points:_ 2

```python
remaining_inventory = ...
    ...
    ...
    ...
remaining_inventory
```

```python
grader.check("task_26")
```

## Submit your Homework to Canvas

Once you have finished working on the homework tasks, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the rubric to know how you will be scored for this assignment.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
3. Select the menu item "File" and "Save Notebook" in the notebook's Toolbar to save your work and create a specific checkpoint in the notebook's work history.
4. Select the menu items "File", "Download" in the notebook's Toolbar to download the notebook (.ipynb) file. 
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded .ipynb file.

**Keep in mind that the autograder does not always check for correctness. Sometimes it just checks for the format of your answer, so passing the autograder for a question does not mean you got the answer correct for that question.**

Tests for task_01 (points: 2)
```
>>> 
>>> import numpy
>>> isinstance(weird_numbers, numpy.ndarray) and len(weird_numbers) == 4
True
```
failure_message: ❌ You didn't provide an array of 4 values. Consider using make_array.
success_message: ✅ You provided an array of 4 values.
```
>>> np.allclose(weird_numbers, np.array([3000.,  2.23606798, -2,  5.04749727]), rtol=1e-03, atol=1e-03) == True
True
```


Tests for task_02 (points: 2)
```
>>> 
>>> import numpy
>>> isinstance(book_title_words, numpy.ndarray) and len(book_title_words) == 3
True
```
failure_message: ❌ You didn't provide an array of 3 values. Consider using make_array.
success_message: ✅ You provided an array of 3 values.
```
>>> 
>>> not any([',' in text for text in book_title_words])
True
```
failure_message: ❌ It looks like you included commas in the text. The three pieces of text in the array should be: 'Eats' 'Shoots' 'Leaves'
success_message: ✅ Great work not including the commas.
```
>>> 
>>> # 
>>> 'and ' in book_title_words.item(2)
True
```
failure_message: ❌ It looks like you didn't include both words in the last piece of text. It should be the actual string: 'and Leaves'.
success_message: ✅ Great work not including 'and' with the last string.


Tests for task_03 (points: 2)
```
>>> 
>>> isinstance(with_commas, str) and isinstance(without_commas, str)
True
```
failure_message: ❌ I doesn't seem that with_commas and without_commas are strings. Make sure to use the join method on a string. Use ', ' for a comma and ' ' for a space.
success_message: ✅ Nice! with_commas and without_commas are both strings.


Tests for task_04 (points: 2)
```
>>> 
>>> third_element != -10
True
```
failure_message: ❌ It looks like you wrote some_numbers.item(3), but the third element has index 2, not 3.
success_message: ✅ Nice work not using the wrong index.


Tests for task_05 (points: 2)
```
>>> 
>>> import datascience
>>> isinstance(elements_of_some_numbers, datascience.Table)
True
```
failure_message: ❌ It doesn't seem that you've assigned a Table to elements_of_some_numbers
success_message: ✅ Nice work assigning a Table to elements_of_some_numbers.


Tests for task_06 (points: 2)
```
>>> 
>>> index_of_last_element <= 145
True
```
failure_message: ❌ It doesn't seem like you've assigned a possible int to index_of_last_element.
success_message: ✅ You assigned a possible int to index_of_last_element.


Tests for task_07 (points: 2)
```
>>> 
>>> isinstance(most_recent_birth_year, int)
True
```
failure_message: ❌ Make sure you are assigning an integer from the president_birth_years array. Try using .item(...)
success_message: ✅ You successfully submitted an integer.


Tests for task_08 (points: 2)
```
>>> 
>>> isinstance(sum_of_birth_years, int)
True
```
failure_message: ❌ Make sure you are assigning a sum of 3 integers from the sum_of_birth_years array.
success_message: ✅ You successfully submitted an integer.


Tests for task_09 (points: 2)
```
>>> 
>>> all(x % 157 == 0 for x in (first_product, second_product, third_product, fourth_product))
True
```
failure_message: ❌ Not all of your products are divisible by 157, so it doesn't seem that you've multiplied the 4 numbers by 157.
success_message: ✅ It seems like you've multiplied the 4 numbers by 157.


Tests for task_10 (points: 2)
```
>>> 
>>> import numpy
>>> isinstance(numbers, numpy.ndarray)
True
```
failure_message: ❌ Make sure you are assigning an array to numbers. Try using make_array()
success_message: ✅ Good work. It seems like you've assigned an array to numbers.


Tests for task_11 (points: 2)
```
>>> 
>>> import numpy as np
>>> isinstance(correct_products, np.ndarray) and all(x % 1577 == 0 for x in correct_products)
True
```
failure_message: ❌ All the values in your array should be divible by 1577.
success_message: ✅ Good work. It seems like all the values in your array are divible by 1577.


Tests for task_12 (points: 2)
```
>>> 
>>> sum(celsius_max_temperatures) != 356705.0
True
```
failure_message: ❌ Looks like you multiplied and subtracted in the wrong order.
success_message: ✅ Nice work. It is easy to multiple and subtract in the wrong order by mistake.
```
>>> 
>>> import numpy as np
>>> celsius_max_temperatures[0] - np.floor(celsius_max_temperatures[0]) == 0 and celsius_max_temperatures[0] != 25
True
```
failure_message: ❌ Make sure you round your results to the nearest integer. Try using np.round.
success_message: ✅ It seems like you've rounded your results to the nearest integer.


Tests for task_13 (points: 2)
```
>>> 
>>> abs(celsius_temperature_ranges[0]) != 12 and celsius_temperature_ranges[0] >= 0
True
```
failure_message: ❌ Make sure that you are not submitting Fahrenheit values.
success_message: ✅ It seems like you are not submitting Fahrenheit values. Nice work!
```
>>> 
>>> import numpy as np
>>> celsius_temperature_ranges[0] - np.floor(celsius_temperature_ranges[0]) != 0
True
```
failure_message: ❌ Make sure to not round your results on this task.
success_message: ✅ It doesn't seem that you rounded the results. Great work catching that detail in the instructions.


Tests for task_14 (points: 2)
```
>>> 
>>> import numpy as np
>>> isinstance(largest_population_change, np.int64)
True
```
failure_message: ❌ It doesn't seem that you are using the numpy functions.
success_message: ✅ Nice work using numpy functions.


Tests for task_15 (points: 2)
```
>>> 
>>> cumulative_sum_answer in {1, 2, 3}
True
```
failure_message: ❌ Make sure to assign 1, 2, or 3 to cumulative_sum_answer.
success_message: ✅ It seems like you've entered a possible answer.


Tests for task_16 (points: 2)
```
>>> 
>>> shortest <= average <= longest
True
```
failure_message: ❌ The average value is between the shortest and the longest values. Make sure to use functions like min, max, and np.average (or np.mean).
success_message: ✅ Your average value is between your shortest and longest values.


Tests for task_17 (points: 2)
```
>>> biggest_decrease != 47
True
```
failure_message: ❌ If you are getting 47 as your answer, you might be computing the biggest change rather than biggest decrease.
success_message: ✅ It seems that you didn't compute the biggest change. Nice work.
```
>>> 
>>> 30 <= biggest_decrease < 47
True
```
failure_message: ❌ The biggest decrease is above 30, but not 47
success_message: ✅ Your value seems to be in a reasonable interval of values.


Tests for task_18 (points: 1)
```
>>> 
>>> difference_from_expected.size == 272
True
```
failure_message: ❌ You don't seem to have the correct number of values. Don't forget to use np.cumsum
success_message: ✅ You have the correct number of values.
```
>>> 
>>> difference_from_expected.item(271) == abs(60 * 272 - sum(waiting_times))
True
```
failure_message: ❌ One of your values doesn't look correct.
success_message: ✅ One of your values looks correct.


Tests for task_19 (points: 2)
```
>>> 
>>> 15 <= average_error <= 25
True
```
failure_message: ❌ Your average_erro value should be at least 15 and at most 25.
success_message: ✅ Your average_error value is in a reasonable interval of numbers.


Tests for task_20 (points: 2)
```
>>> 
>>> import datascience
>>> isinstance(fruits, datascience.Table)
True
```
failure_message: ❌ Make sure assign a Table to the name fruits.
success_message: ✅ You've assigned a Table to fruits.


Tests for task_21 (points: 1)
```
>>> 
>>> inventory.sort(0).column(0).item(0) == 25274
True
```
failure_message: ❌ One of your values doesn't look correct.
success_message: ✅ One of your values looks correct.


Tests for task_22 (points: 2)
```
>>> 
>>> all_different in {True, False}
True
```
failure_message: ❌ Make sure to respond with either True or False.
success_message: ✅ You've provided a resopnse of True or False.


Tests for task_23 (points: 1)
```
>>> 
>>> sales.sort(0).column(0).item(1) == 26187
True
```
failure_message: ❌ One of your values doesn't look correct.
success_message: ✅ One of your values looks correct.


Tests for task_24 (points: 2)
```
>>> 
>>> total_fruits_sold > 10
True
```
failure_message: ❌ We're asking for the number of *pieces of fruit, not the number of kinds of fruit or the number of boxes from which there were sales.
success_message: ✅ Your value for total_fruits_sold seems reasonable.


Tests for task_25 (points: 2)
```
>>> 
>>> 50 <= total_revenue <= 150
True
```
failure_message: ❌ As a hint, you want to multiply the count sold in each box by the per-item price of fruits in that box. You can use element-wise multiplication for that. Then you want the sum of those products. Use sum().
success_message: ✅ Your value for total_revenue is in a reasonable interval of values.


Tests for task_26 (points: 2)
```
>>> 
>>> remaining_inventory.num_columns == 3
True
```
failure_message: ❌ It looks like your table doesn't have all 3 columns that are in the inventory table.
success_message: ✅ You have the correct number of columns.
```
>>> 
>>> remaining_inventory.column("count").item(0) != 45
True
```
failure_message: ❌ It looks like you forgot to subtract off the sales.
success_message: ✅ It seems like you subtractted off the sales.
