```python
# Initialize Otter
import otter
grader = otter.Notebook("hw06.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

<div style="text-align: center;">
    <h1>Homework 06: Probability and Simulation</h1>
    <em>View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.</em>
</div>

**Reading**: 
* [Randomness](https://inferentialthinking.com/chapters/09/Randomness.html) 
* [Sampling and Empirical Distributions](https://inferentialthinking.com/chapters/10/Sampling_and_Empirical_Distributions.html)
* [Testing Hypotheses](https://inferentialthinking.com/chapters/11/Testing_Hypotheses.html)

Please complete this notebook by filling in the cells provided. 

For all problems that you must write our explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on.

Run the following cell:

```python
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## Roulette

A Nevada roulette wheel has 38 pockets and a small ball that rests on the wheel. When the wheel is spun, the ball comes to rest in one of the 38 pockets. That pocket is declared the winner. 

The pockets are labeled 0, 00, 1, 2, 3, 4, ... , 36. Pockets 0 and 00 are green, and the other pockets are alternately red and black. The table `wheel` is a representation of a Nevada roulette wheel. **Note that *both* columns consist of strings.** Below is an example of a roulette wheel!

<img src="./roulette_wheel.jpeg" width="330px">

Run the cell below to load the `wheel` table.

```python
wheel = Table.read_table('roulette_wheel.csv', dtype=str)
wheel
```

### Betting on Red

If you bet on *red*, you are betting that the winning pocket will be red. This bet *pays 1 to 1*. That means if you place a one-dollar bet on red, then:

- If the winning pocket is red, you gain 1 dollar. That is, you get your original dollar back, plus one more dollar.
- if the winning pocket is not red, you lose your dollar. In other words, you gain -1 dollars.

Let's see if you can make money by betting on red at roulette.

#### Task 01 📍

Define a function `dollar_bet_on_red`. The function definition should:
1. Have one argument `color` that is a `str` for the name of a color.
2. Return your gain in dollars as an `int` if that color had won and you had placed a one-dollar bet on red. 

Consider the following as you work:
* Remember that the gain can be negative.
* Make sure your function returns an integer.
* You can assume that the only colors that will be passed as arguments are `'red'`, `'black'`, and `'green'`. Your function doesn't have to check that the input is correct.

_Points:_ 3

```python
def dollar_bet_on_red(...):
    ...
```

```python
grader.check("task_01")
```

Run the cell below to make sure your function is working.

```python
print(dollar_bet_on_red('green'))
print(dollar_bet_on_red('black'))
print(dollar_bet_on_red('red'))
```

#### Task 02 📍

1. Add a column labeled `'Winnings: Red'` as the last column in the table `wheel` such that, for each pocket, the column should contain your gain in dollars if that pocket won and you had bet one dollar on red. 
2. Your code should apply the function `dollar_bet_on_red` on the `wheel` table to create an array `red_winnings`.

_Points:_ 3

```python
red_winnings = ...
wheel = ...
wheel
```

```python
grader.check("task_02")
```

### Simulating 10 Bets on Red

Roulette wheels are set up so that each time they are spun, the winning pocket is equally likely to be any of the 38 pockets regardless of the results of all other spins. Let's see what would happen if we decided to bet one dollar on red each round.

#### Task 03 📍

1. Create a table `ten_bets` by sampling the table `wheel` to simulate 10 spins of the roulette wheel. 
2. Your table should have the same three column labels as in `wheel`. 
3. Once you've created that table, set `sum_bets` to your net gain, an integer, in all 10 bets, assuming that you bet one dollar on red each time.

While you are working on this, it may be helpful to print out `ten_bets` after you create it!

_Points:_ 5

```python
ten_bets = ...
sum_bets = ...(ten_bets.column('Winnings: Red'))
sum_bets
```

```python
grader.check("task_03")
```

Run the cells above a few times to see how much money you would make if you made 10 one-dollar bets on red. Making a negative amount of money doesn't feel good, but it is a reality in gambling. Casinos are a business, and they make money when gamblers lose.

#### Task 04 📍

Let's see what would happen if you made more bets. 

Define a function `net_gain_red` with one arguement that takes the number (`int`) of bets and returns the net gain (`int`) in that number of one-dollar bets on red.

Reference the `wheel` table that you've defined in the notebook within your function.

_Points:_ 2

```python
def net_gain_red(...):
    ...
    
net_gain_red(10)
```

```python
grader.check("task_04")
```

#### Task 05 📍

1. Complete the cell below to simulate the net gain in 200 one-dollar bets on red, repeating the process 10,000 times. 
2. After the cell is run, `all_gains_red` should be an array with 10,000 entries, each of which is the net gain in 200 one-dollar bets on red.

_Points:_ 2

```python
num_bets = ...
repetitions = ...

all_gains_red = ...
...

len(all_gains_red) # Do not change this line! Check that all_gains_red is length 10000.
```

```python
grader.check("task_05")
```

Run the cell below to visualize the results of your simulation.

```python
gains = Table().with_columns('Net Gain on Red', all_gains_red)
gains.hist(bins = np.arange(-80, 41, 4))
```

#### Task 06 📍

Using the histogram above, decide whether the following statement is true or false:

>If you make 200 one-dollar bets on red, your chance of losing money is more than 50%.

Assign `loss_more_than_50` to either `True` or `False` depending on your answer to the question.

_Points:_ 2

```python
loss_more_than_50 = ...
```

```python
grader.check("task_06")
```

### Betting on a Split

If betting on red doesn't seem like a good idea, maybe a gambler might want to try a different bet. A bet on a *split* is a bet on two consecutive numbers such as 5 and 6. This bets pays 17 to 1. That means if you place a one-dollar bet on the split 5 and 6, then:

- If the winning pocket is either 5 or 6, your gain is 17 dollars.
- If any other pocket wins, you lose your dollar, so your gain is -1 dollars.

#### Task 07 📍

Define a function `dollar_bet_on_5_6_split`.
1. The function should have 1 argument (type `str`) that represents the pocket number
2. The function should return the gain in dollars if that pocket won and they had bet one dollar on the 5-6 split.

Remember that the pockets are represented as strings.

_Points:_ 4

```python
def dollar_bet_on_5_6_split(...):
    ...
```

```python
grader.check("task_07")
```

Run the cell below to check that your function is doing what it should.

```python
print(dollar_bet_on_5_6_split('5'))
print(dollar_bet_on_5_6_split('6'))
print(dollar_bet_on_5_6_split('00'))
print(dollar_bet_on_5_6_split('23'))
```

#### Task 08 📍

Add a column `'Winnings: 5-6 Split'` to the end of the `wheel` table. For each pocket, the column should contain your gain in dollars if that pocket won and you had bet one dollar on the 5-6 split.

_Points:_ 3

```python
split_winnings = ...
wheel = ...
```

```python
grader.check("task_08")
```

#### Task 09 📍

1. Simulate the net gain in 200 one-dollar bets on the 5-6 split.
2. Repeat the simulation for a total of 10,000 times.
3. Store your gains in the array `all_gains_split`.

_Points:_ 2

```python
all_gains_split = ...
...

# Do not change the two lines below
gains = gains.with_columns('Net Gain on Split', all_gains_split)
gains.hist(bins = np.arange(-200, 150, 20))
```

```python
grader.check("task_09")
```

#### Task 10 📍

Look carefully at the visualization above, and assign `histogram_statements` to an array of the numbers of each statement below that can be correctly inferred from the overlaid histogram.

1. If you bet one dollar 200 times on a split, your chance of losing money is more than 50%.
2. If you bet one dollar 200 times in roulette, your chance of making more than 50 dollars is greater if you bet on a split each time than if you bet on red each time.
3. If you bet one dollar 200 times in roulette, your chance of losing more than 50 dollars is greater if you bet on a split each time than if you bet on red each time.

Notice that you've already seen one of these statements in a prior question.

_Points:_ 3

```python
histogram_statements = ...
```

```python
grader.check("task_10")
```

## Chances

Before you do this exercise, make sure you understand the logic behind all the examples in [Section 9.5](https://inferentialthinking.com/chapters/09/5/Finding_Probabilities.html). 

Good ways to approach probability calculations include:

- Thinking one trial at a time: What does the first one have to be? Then what does the next one have to be?
- Breaking up the event into distinct ways in which it can happen.
- Seeing if it is easier to find the chance that the event does not happen.

### Finding Chances

On each spin of a roulette wheel, all 38 pockets are equally likely to be the winner regardless of the results of other spins. Among the 38 pockets, 18 are red, 18 black, and 2 green. In each part below, write an expression that Python evaluates to the chance of the event described.

#### Task 11 📍

The winning pocket is black on all of the first three spins.

_Points:_ 2

```python
first_three_black = ...
```

```python
grader.check("task_11")
```

#### Task 12 📍

The color green never wins in the first 10 spins.

_Points:_ 2

```python
no_green = ...
```

```python
grader.check("task_12")
```

#### Task 13 📍

The color green wins **at least** once on the first 10 spins.

_Points:_ 2

```python
at_least_one_green = ...
```

```python
grader.check("task_13")
```

#### Task 14 📍

One of the three colors **never** wins in the first 10 spins.

_Points:_ 2

```python
lone_winners = ...
```

```python
grader.check("task_14")
```

### Comparing Chances

In each of the following 3 questions, two events A and B are described. Choose from one of the following three options and set each answer variable to a single integer:

1. Event A is more likely than Event B
2. Event B is more likely than Event A
3. The two events have the same chance.

You should be able to make the choices **without calculation**. Good ways to approach this exercise include imagining carrying out the chance experiments yourself, one trial at a time, and by thinking about the [law of averages](https://inferentialthinking.com/chapters/10/1/Empirical_Distributions.html#the-law-of-averages).

#### Task 15 📍

A child picks four times at random from a box that has four toy animals: a bear, an elephant, a giraffe, and a kangaroo.

- Event A: all four different animals are picked (assuming the child picks without replacement)
- Event B: all four different animals are picked (assuming the child picks with replacement)

_Points:_ 2

```python
toys_option = ...
```

```python
grader.check("task_15")
```

#### Task 16 📍

In a lottery, two numbers are drawn at random without replacement from the integers 1 through 1000.

- Event A: The number 8 is picked on both draws
- Event B: The same number is picked on both draws


_Points:_ 2

```python
lottery_option = ...
```

```python
grader.check("task_16")
```

#### Task 17 📍

A fair coin is tossed repeatedly.

- Event A: There are 600 or more heads in 1000 tosses
- Event B: There are 60 or more heads in 100 tosses

Remember the law of averages!

_Points:_ 2

```python
coin_option = ...
```

```python
grader.check("task_17")
```

## Submit your Homework to Canvas

Once you have finished working on the homework tasks, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the rubric to know how you will be scored for this assignment.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
3. Select the menu item "File" and "Save Notebook" in the notebook's Toolbar to save your work and create a specific checkpoint in the notebook's work history.
4. Select the menu items "File", "Download" in the notebook's Toolbar to download the notebook (.ipynb) file. 
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded .ipynb file.

**Keep in mind that the autograder does not always check for correctness. Sometimes it just checks for the format of your answer, so passing the autograder for a question does not mean you got the answer correct for that question.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 3)
```
>>> 
>>> dollar_bet_on_red('red') == 1
True
```
failure_message: ❌ Your function should return 1 when it is called on 'red'.
success_message: ✅ Your function correctly returns 1 when it is called on 'red'.


Tests for task_02 (points: 3)
```
>>> 
>>> isinstance(red_winnings, np.ndarray)
True
```
failure_message: ❌ red_winnings should be an array.
success_message: ✅ It seems that red_winnings is an array.
```
>>> 
>>> wheel.labels[2] == "Winnings: Red"
True
```
failure_message: ❌ wheel is not a table with the correct labels.
success_message: ✅ wheel is a table with the correct labels.


Tests for task_03 (points: 5)
```
>>> 
>>> ten_bets.num_rows == 10
True
```
failure_message: ❌ ten_bets should have 10 rows.
success_message: ✅ ten_bets has 10 rows.
```
>>> 
>>> set(["Pocket", "Color", "Winnings: Red"]) == set(ten_bets.labels)
True
```
failure_message: ❌ ten_bets should have the same column labels as wheels.
success_message: ✅ ten_bets has the same column labels as wheels.
```
>>> 
>>> isinstance(sum_bets, int) or isinstance(sum_bets, np.int64)
True
```
failure_message: ❌ sum_bets should be an integer.
success_message: ✅ sum_bets is an integer.


Tests for task_04 (points: 2)
```
>>> 
>>> net_gain_red(10000) != net_gain_red(10000)
True
```
failure_message: ❌ There is essentially no chance that the net gain for 2 rounds of 10,000 bets should the same.
success_message: ✅ There is essentially no chance that the net gain for 2 rounds of 10,000 bets should the same.


Tests for task_05 (points: 2)
```
>>> 
>>> len(all_gains_red) == 10000
True
```
failure_message: ❌ all_gains_red should be an array with 10,000 items.
success_message: ✅ all_gains_red is an array with 10,000 items.


Tests for task_06 (points: 2)
```
>>> 
>>> isinstance(loss_more_than_50, bool)
True
```
failure_message: ❌ loss_more_than_50 should be a bool value.
success_message: ✅ loss_more_than_50 is a bool value.


Tests for task_07 (points: 4)
```
>>> 
>>> dollar_bet_on_5_6_split('5') == 17
True
```
failure_message: ❌ Your function didn't return the correct value for the 5 pocket.
success_message: ✅ Your function returned the correct value for the 5 pocket.
```
>>> 
>>> dollar_bet_on_5_6_split('00') == -1
True
```
failure_message: ❌ Your function didn't return the correct value for the 00 pocket.
success_message: ✅ Your function returned the correct value for the 00 pocket.


Tests for task_08 (points: 3)
```
>>> 
>>> wheel.labels[3] == 'Winnings: 5-6 Split'
True
```
failure_message: ❌ Your last column label should be 'Winnings: 5-6 Split'.
success_message: ✅ Your last column label is 'Winnings: 5-6 Split'.
```
>>> 
>>> wheel.where('Pocket', '5').column('Winnings: 5-6 Split').item(0) == 17
True
```
failure_message: ❌ The added column value associated with the 5 pocket is not correct.
success_message: ✅ The added column value associated with the 5 pocket is correct.


Tests for task_09 (points: 2)
```
>>> 
>>> len(all_gains_split) == 10000
True
```
failure_message: ❌ There should be 10,000 items in all_gains_split.
success_message: ✅ There are 10,000 items in all_gains_split.


Tests for task_10 (points: 3)
```
>>> 
>>> isinstance(histogram_statements, np.ndarray)
True
```
failure_message: ❌ histogram_statements should be an array.
success_message: ✅ histogram_statements is an array.
```
>>> 
>>> isinstance(histogram_statements.item(0), int)
True
```
failure_message: ❌ histogram_statements should contain at least one integer.
success_message: ✅ histogram_statements contains at least one integer.


Tests for task_11 (points: 2)
```
>>> 
>>> 0 <= first_three_black <= 1
True
```
failure_message: ❌ first_three_black should be a numerical value between 0 and 1, inclusive.
success_message: ✅ first_three_black is a numerical value between 0 and 1, inclusive.


Tests for task_12 (points: 2)
```
>>> 
>>> 0 <= no_green <= 1
True
```
failure_message: ❌ no_green should be a numerical value between 0 and 1, inclusive.
success_message: ✅ no_green is a numerical value between 0 and 1, inclusive.


Tests for task_13 (points: 2)
```
>>> 
>>> 0 <= at_least_one_green <= 1
True
```
failure_message: ❌ at_least_one_green should be a numerical value between 0 and 1, inclusive.
success_message: ✅ at_least_one_green is a numerical value between 0 and 1, inclusive.


Tests for task_14 (points: 2)
```
>>> 
>>> 0 <= lone_winners <= 1
True
```
failure_message: ❌ lone_winners should be a numerical value between 0 and 1, inclusive.
success_message: ✅ lone_winners is a numerical value between 0 and 1, inclusive.


Tests for task_15 (points: 2)
```
>>> 
>>> toys_option in [1, 2, 3]
True
```
failure_message: ❌ 1, 2, or 3 should be assigned to toys_option.
success_message: ✅ This is a possible value for toys_option.


Tests for task_16 (points: 2)
```
>>> 
>>> lottery_option in [1, 2, 3]
True
```
failure_message: ❌ 1, 2, or 3 should be assigned to lottery_option.
success_message: ✅ This is a possible value for lottery_option.


Tests for task_17 (points: 2)
```
>>> 
>>> coin_option in [1, 2, 3]
True
```
failure_message: ❌ 1, 2, or 3 should be assigned to coin_option.
success_message: ✅ This is a possible value for coin_option
