```python
# Initialize Otter
import otter
grader = otter.Notebook("lab05.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

<div style="text-align: center;">
    <h1>Lab 05: Visualizations</h1>
    <em>View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.</em>
</div>

This week, you will focus on creating visualizations to support the analysis of BART ridership data in addition to practicing more with the `group` and `join` methods.

First, set up the notebook by running the cell below.

```python
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
```

## SF311

The City and County of San Francisco established SF311 to provide an easy-to-remember telephone number that connects residents, businesses, and visitors to Customer Service Representatives ready to help with general government information and services in 2007. SF311 started taking web requests in 2008, social media in 2009, and mobile requests in 2010 from apps that connect to the open311 API. Departments can also use the 311 system to report inter-agency issues.

This dataset includes cases generally associated with a place or thing (for example parks, streets, or buildings) and created July 1, 2008 or later. Cases generally logged by a user regarding their own needs (for example, property or business tax questions, parking permit requests) are not included. Cases prior to July 1, 2008 are not included due to significant data structure changes in the spring of 2008. General Requests are multipurpose requests that can be routed to the majority of City Agencies. SF311 case data is provided as is and without warranty but works with its departmental partners to continuously improve business processes and data quality.

The following is a list of the included case types:

* 311 External Request
* Abandoned Vehicle
* Blocked Street or Sidewalk
* Catch Basin Maintenance
* Color Curb
* Construction Zone Permits
* Damaged Property
* DPW Volunteer Programs
* Encampments
* Entertainment Commission
* General Request (Multiple Agencies)
* Graffiti
* Homeless Concerns
* Illegal Postings
* Litter Receptacles
* MUNI Feedback
* Noise Report
* Parking Enforcement
* Rec and Park Requests
* Residential Building Request
* Sewer Issues
* SFHA Requests
* Sidewalk or Curb
* Sign Repair
* Street and Sidewalk Cleaning
* Street Defects
* Streetlights
* Temporary Sign Request
* Tree Maintenance

Run the following code cell to to load the current (February 28, 2023) 311 data sourced from [the DataSF 311 Cases data set](https://data.sfgov.org/City-Infrastructure/311-Cases/vw6y-z8j6).

```python
# You are not response for the following code. Lambda functions are covered in more advanced Python courses
sf311 = Table.read_table('sf311.csv')
sf311 = sf311.with_column('supervisor_district', sf311.apply(lambda x: str(int(x)), 'supervisor_district'))
sf311
```

The data contains information on the supervisor districts in the city. These districts are regions of the city that are represented by an elected district supervisor. The districts were re-drawn in 2022. You can read more about the districts and the 2022 changes in the SF Chronicle article [_These 11 maps show exactly how San Francisco's supervisor districts will be redrawn_](https://www.sfchronicle.com/projects/2022/san-francisco-redistricting-final-map/)

### Join

The `join` table method allows you to combine the information in one table with the information in another table by relating the data through matching values in a specific column from each table.

For example, the following two tables `tbl1` and `tbl2` share common information in columns `'tbl1_col1'` and `'tbl2_col2'`.

```python
tbl1 = Table().with_columns('tbl1_col1', ['A', 'B', 'A'], 'tbl1_col2', [1, 3, 1], 'tbl1_col3', ['Yes', 'Maybe', 'No'])
tbl2 = Table().with_columns('tbl2_col1', [True, False, False], 'tbl2_col2', ['A', 'B', 'C'])
display(tbl1)
display(tbl2)
```

The two tables can be joined using the following code:

```python
tbl2.join('tbl2_col2', tbl1, 'tbl1_col1')
```

Here are a few things to notice about the resulting table:
* The first column is the column from `tbl2` that was used to join the data in the two tables, the second column is the rest of the column data from `tbl2`, and the remaining columns are the columns in `tbl1`, except for the column used to join the data in the the two tables.
* There was not a match for for `'C'` in `tbl1`, so there is not a row in in the joined table containing information related to `'C'`.

### Task 01 📍

There can be several reasons why you might want to use a numerical code for categorical data. There are many tools that work on numerical data types that might be helpful for analyzing categorical data. The following table show a numerical translation for the descriptions `'Open'` and `'Closed'` as `1` and `0`, respectively.

```python
code_table = Table().with_columns(
    'status_description', ['Open', 'Closed'],
    'status_code', [1, 0]
)
code_table
```

* Use the `join` table method to create a table called `sf311_coded_status` with same information as `sf311` based on whether or not the case is open. 
* `sf311_coded_status` should only have the columns, in the following order, `'status_code'`, `'requested_datetime'`, `'closed_date'`, `'service_name'`, and `'supervisor_district'`

```python
sf311_coded_status = ...
sf311_coded_status
```

```python
grader.check("task_01")
```

### Group

The `group` table method with a single argument counts the number of rows for each category in a column. The result contains one row per unique value in the grouped column.

### Task 02 📍

1. Use the `group` method to create a table called `sf311_by_service` with two columns (`'service name'` and `'count'`) showing the a row for each unique service name and the number of time that service request appears in the data set. 
2. Sort the rows in the table such that the services that appear the most often in the dataset appear at the top.

Note: 311 External Requests have been removed form the data, so you won't find information on that service.

```python
sf311_by_service = ...
sf311_by_service
```

```python
grader.check("task_02")
```

### Task 03 📍🔎

Use the `sf311_by_service` table to create a visualization showing the distribution of services in the 311 data. Check with a classmate, the instructor, or a tutor to see if your visualization looks good.

```python
...
```

The optional second argument of `group` names the function that will be used to aggregate values in other columns for all of those rows. For example, the following code applies NumPy\'s average function to the status codes to provide the proportion of open cases for each service name. (The average of a collection of 1\'s and 0\'s values is the proportion of 1\'s in the collection.)

The `group` method will try and apply `np.average` to every column in the table, except for the one used to group the data by. Notice that we first reduced the table to just the column information we want to group by and the column we want to apply `np.average` to. 

```python
sf311_coded_status_reduced_to_2 = sf311_coded_status.select('service_name', 'status_code')
sf311_coded_status_reduced_to_2.group('service_name', np.average)
```

It seems that 87.5% of the damaged property 311 cases are still open.

The group method also allows us to classify individuals according to multiple variables (cross-classify). Instead of providing a column label or column index for the first argument to `group`, you can provide a list of column labels or column indices to group by. Each row in the resulting table will show a unique combination of the values associated with the provided columns.

```python
sf311.group(['service_name', 'supervisor_district']).sort('count', True)
```

### Task 04 📍

Create a table called `sf311_coded_status_relabled` using the following steps:
1. Reduce the `sf311_coded_status` table to only have the columns `'service_name'`, `'supervisor_district'`, and `'status_code'`.
2. Use the `group` method to cross-classify the data based on the values in the `'service_name'` and `'supervisor_district'` columns, and use `np.average` for the the optional second parameter for `group` to include the proportion of open cases for each service name and supervisor district combination.
3. Relabel the status code average column to be `'proportion_open'`.

For the auto-grader, make sure you use `np.average` when you group the data.

```python
sf311_coded_status_reduced_to_3 = ...
sf311_coded_status_by_service_and_district = ...
sf311_coded_status_relabled = ...
sf311_coded_status_relabled
```

```python
grader.check("task_04")
```

### Task 05 📍

Focusing on just the supervisor districts and the case status, use some of the tools you have learned to determine the supervisor district with the largest total proportion of open cases. 

1. Create the table `sf311_district_prop_open` using the `group` method and the `np.average` function that has two columns, `'supervisor_district'` and `'Proportion Open'` showing the proportion of open cases for each district.
2. Assign `district_with_largest_prop_open` to the district number (`str`) with the largest proportion of open cases.

```python
sf311_district_prop_open = ...
district_with_largest_prop_open = ...
district_with_largest_prop_open
```

```python
grader.check("task_05")
```

Next, focus on the closed 311 cases. That information has been in a table named `sf311_closed`.

```python
sf311_closed = sf311.where('status_description', 'Closed').drop('status_description')
sf311_closed
```

### Datetime Data

The data in the `'requested_datetime'` and '`closed_date'` columns of `sf311_closed` are in format that you might not be familiar with, but it (ISO format) is a standard format. Python has tools to work with that data, and the related module for that are curious is called [`datetime`](https://docs.python.org/3/library/datetime.html). 

**You will not be tested on the information in the `datetime` module.**

The following function returns the number of seconds it took for SF311 to close each 311 case in the data set.

```python
def completion_time(requested_datetime, closed_date):
    '''Returns the amount of time (in seconds) from the requested date to the closed date for the 311 case.'''
    from datetime import datetime
    time_delta = datetime.fromisoformat(closed_date) - datetime.fromisoformat(requested_datetime)
    return time_delta.seconds
```

The following code shows you how the function was used to determine that first case listed in `sf311_closed` took 1,200 seconds to close.

```python
completion_time('2023-02-27T23:48:00.000', '2023-02-28T00:08:00.000')
```

### Task 06 📍

Use the `completion_time` function to update the `sf311_closed` table with an additional column `'Completion Time'` that contains the time (in seconds) it took for the case to be closed.

```python
completion_times = ...
sf311_closed = ...
sf311_closed
```

```python
grader.check("task_06")
```

### Task 07 📍🔎

<!-- BEGIN QUESTION -->

Create a histogram showing the distribution of case completion times. Use the `unit='second'` parameter. Check with a classmate, the instructor, or a tutor to see if your visualization looks good.

_Type your answer here, replacing this text._

```python
...

# Leave the following code to show the plot title.
plots.title('311 Case Completion Times')
plots.show()
```

<!-- END QUESTION -->

### Task 08 📍

Assign `completion_time_hist_description` to one of the following strings that best describes the histogram: `'skewed_left'`, `'skewed_right'`, '`neither_skewed_left_or_skewed_right'`.

```python
completion_time_hist_description = ...
completion_time_hist_description
```

```python
grader.check("task_08")
```

### Task 09 📍

The following table shows you the number of 311 cases within each bar of the above histogram. For example, there are 43 data values included in the bin cover the completion times from 3034 seconds up to, but not including, 6027 seconds.

```python
sf311_closed.bin('Completion Time').show()
```

Suppose that the bins were created such that there was a bin $[3034, 9020)$ that combined the information from $[3034, 6027)$ and $[6027, 9020)$. Assign the height that wider bar to the name `wide_bar_height`. Make sure to use the intermediate variables to guide your calculation.

```python
percent_of_data_in_wide_bin = ...
wide_bar_height = ...
wide_bar_height
```

```python
grader.check("task_09")
```

### Task 10 📍🔎

<!-- BEGIN QUESTION -->

Early you created the table `sf311_district_prop_open` that contained the proportion of open 311 cases for each district. Now, you have the table `average_completion_time_by_district` that contains the average completion time for a 311 case for each district.

1. Create a visualization to explore whether or not there is a strong association between these two variables. 
2. Provide a sentence with your interpretation of the scatter plot in terms of whether or not it provides visual evidence of a strong association between the proportion of open cases in a district and the average case completion time in that district.

Check with a classmate, the instructor, or a tutor to see if your visualization and interpretation look good.


_Type your answer here, replacing this text._

```python
...
```

<!-- END QUESTION -->

Great work so far! There is a lot to analyze in this data set. If you pay taxes in San Francisco, I encourage you to dig deeper into this data set since you pay for this service and data! For now, make sure to submit your work. 

## Submit your Lab to Canvas

Once you have finished working on the lab questions, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the requirements for a Complete score for this lab assignments.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
4. Select the menu items "File", "Save and Export Notebook As...", and "HTML (.html)" in the notebook's Toolbar to download an HTML version of this notebook file.
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded HTML file.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: None)
```
>>> 
>>> isinstance(sf311_coded_status, Table)
True
```
failure_message: ❌ sf311_coded_status should be a Table. Make sure you are using the join method.
success_message: ✅ sf311_coded_status is a Table.
```
>>> 
>>> (sf311_coded_status.labels) == ('status_code', 
...                                 'requested_datetime',
...                                 'closed_date',
...                                 'service_name',
...                                 'supervisor_district')
True
```
failure_message: ❌ sf311_coded_status should have exactly the same labels and in the same ordered mentioned. You might need to drop the 'status_description' column or reorder the columns.
success_message: ✅ sf311_coded_status has the correct labels that are in the correct order.
```
>>> 
>>> print(sf311_coded_status.sort('requested_datetime'))
status_code | requested_datetime      | closed_date             | service_name                 | supervisor_district
0           | 2023-02-27T13:46:00.000 | 2023-02-27T19:21:47.000 | Parking Enforcement          | 6
1           | 2023-02-27T13:46:00.000 | nan                     | Rec and Park Requests        | 5
1           | 2023-02-27T13:46:37.000 | nan                     | Sign Repair                  | 4
1           | 2023-02-27T13:46:53.000 | nan                     | Sewer Issues                 | 1
0           | 2023-02-27T13:47:00.000 | 2023-02-27T14:08:33.000 | Parking Enforcement          | 6
1           | 2023-02-27T13:47:39.000 | nan                     | Streetlights                 | 4
0           | 2023-02-27T13:47:52.000 | 2023-02-27T16:49:49.000 | Sewer Issues                 | 4
1           | 2023-02-27T13:47:55.000 | nan                     | Sign Repair                  | 5
1           | 2023-02-27T13:48:25.000 | nan                     | Sign Repair                  | 5
0           | 2023-02-27T13:48:59.000 | 2023-02-27T19:45:36.000 | Street and Sidewalk Cleaning | 10
... (979 rows omitted)
```
failure_message: ❌ sf311_coded_status doesn't seem correct. Double check with someone about how you are using the join method.
success_message: ✅ sf311_coded_status seems to be correct.


Tests for task_02 (points: None)
```
>>> 
>>> isinstance(sf311_by_service, Table)
True
```
failure_message: ❌ sf311_by_service should be a Table. Use sf311.group() to complete this task.
success_message: ✅ sf311_by_service is a Table
```
>>> 
>>> sf311_by_service.num_rows == 28
True
```
failure_message: ❌ sf311_by_service should have 28 rows. Make sure to group by the service_name column in the sf311 table.
success_message: ✅ sf311_by_service should have 28 rows
```
>>> 
>>> set(sf311_by_service.labels) == {'service_name', 'count'}
True
```
failure_message: ❌ sf311_by_service should have a 'service_name' column and a 'count' column. Make sure to group by the service_name column in the sf311 table.
success_message: ✅ sf311_by_service has the correct column labels.


Tests for task_04 (points: None)
```
>>> 
>>> sf311_coded_status_relabled.num_columns == 3
True
```
failure_message: ❌ sf311_coded_status_relabled doesn't seem to have 3 columns. Make sure you selected the columns above.
success_message: ✅ sf311_coded_status_relabled has 3 columns
```
>>> 
>>> np.isclose(sum(sf311_coded_status_relabled.drop('service_name', 'supervisor_district').column(0)), 127.53924096391582)
True
```
failure_message: ❌ sf311_coded_status_relabled doesn't to contain the correct proportion of open cases for each service name and district combinations. Make sure you are using the np.average function with group method.
success_message: ✅ sf311_coded_status_relabled seems to contain the correct proportion of open cases for each service name and district combinations.
```
>>> 
>>> 'proportion_open' in sf311_coded_status_relabled.labels
True
```
failure_message: ❌ sf311_coded_status_relabled doesn't seem to have a column labeled 'proportion_open'. Make sure to relabeled 'status_code average'.
success_message: ✅ sf311_coded_status_relabled has a column labeled 'proportion_open'
```
>>> 
>>> sf311_coded_status_relabled.where('service_name', 'Encampments').where('supervisor_district', '2').column('proportion_open').item(0) == 0.875
True
```
failure_message: ❌ The value we checked in sf311_coded_status_relabled doesn't seem correct.
success_message: ✅ The value we checked in sf311_coded_status_relabled seems correct.


Tests for task_05 (points: None)
```
>>> 
>>> 'Proportion Open' in sf311_district_prop_open.labels
True
```
failure_message: ❌ sf311_district_prop_open should contain a column called 'Proportion Open'. You will probably need to relabel a column in your sf311_district_prop_open table.
success_message: ✅ sf311_district_prop_open contains a column called 'Proportion Open'.
```
>>> 
>>> np.isclose(np.sum(sf311_district_prop_open.column('Proportion Open')), 7.7674900296331648)
True
```
failure_message: ❌ sf311_district_prop_open doesn't seem to contain the correct proportions of open cases. Check with someone to check your thought process.
success_message: ✅ sf311_district_prop_open seems to contain the correct proportions of open cases.
```
>>> 
>>> isinstance(district_with_largest_prop_open, str)
True
```
failure_message: ❌ district_with_largest_prop_open should represent a string.
success_message: ✅ district_with_largest_prop_open represents a string.
```
>>> 
>>> district_with_largest_prop_open == '7'
True
```
failure_message: ❌ You haven't correctly identified the district with the largest proportion of open 311 cases.
success_message: ✅ You have correctly identified the district with the largest proportion of open 311 cases.


Tests for task_06 (points: None)
```
>>> 
>>> sf311_closed.num_columns == 5 and 'Completion Time' in sf311_closed.labels
True
```
failure_message: ❌ sf311_closed should have 5 columns, including 'Completion Time'. You can use the with_column method to add a column to an extisting table.
success_message: ✅ sf311_closed has 5 columns, include 'Completion Time'
```
>>> 
>>> np.sum(sf311_closed.column('Completion Time')) == 2413824
True
```
failure_message: ❌ The time to time_to_complete values don't seem correct. Did you use the apply method with the completion_time function? Make sure you have the column labels in the correct order with the apply method.
success_message: ✅ The time_to_complete values seem correct'


Tests for task_08 (points: None)
```
>>> 
>>> completion_time_hist_description == 'skewed_right'
True
```
failure_message: ❌ Since the histogram tapers off to the right, this histogram could be generally desribed as oen that is skewed right.
success_message: ✅ Since the histogram tapers off to the right, this histogram could be generally desribed as oen that is skewed right.


Tests for task_09 (points: None)
```
>>> 
>>> 0 <= percent_of_data_in_wide_bin <= 100 and np.round(percent_of_data_in_wide_bin, 3) != 0.261
True
```
failure_message: ❌ percent_of_data_in_wide_bin should reflect a percentage (between 0 and 100). Make sure you are calculating the number of cases included in the new bin divdied by the total number of cases. Also, don't foget to conver the proportion to a percentage value.
success_message: ✅ percent_of_data_in_wide_bin seems to reflect the correct percentage.
```
>>> 
>>> np.isclose(2 ** wide_bar_height, 1.0030271553154249)
True
```
failure_message: ❌ Your value for wide_bar_height doesn't seem correct. Make sure you are divding the percentage of the data in the new bin by the width of the new bin.
success_message: ✅ It seems that you have the corret value for wide_bar_height.
