```python
# Initialize Otter
import otter
grader = otter.Notebook("lab08.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

<div style="text-align: center;">
    <h1>Lab 08: A/B Testing</h1>
    <em>View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.</em>
</div>

This week's lab will focus on A/B Testing using data from the ever-popular British television show, [*The Great British Bakeoff*](https://en.wikipedia.org/wiki/The_Great_British_Bake_Off).

```python
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
from matplotlib import patches
from ipywidgets import interact, interactive, fixed
import ipywidgets as widgets
```

## The Great British Bake Off

>"The Great British Bake Off (often abbreviated to Bake Off or GBBO) is a British television baking competition, produced by Love Productions, in which a group of amateur bakers compete against each other in a series of rounds, attempting to impress a group of judges with their baking skills" [Source: Wikipedia](https://en.wikipedia.org/wiki/The_Great_British_Bake_Off)

For every week of the competition, the judges assign one contestant the title "Star Baker". Ultimately, one winner is crowned every season. Using this information, we would like to investigate how winning Star Baker awards affects the odds of winning a season of the show.

## Task 01 📍🔎

<!-- BEGIN QUESTION -->

We want to know whether winning more Star Baker awards causes a change in likelihood of winning the season. Why is it not sufficient to compare star baker rates for winners and losers?

_Check your response with someone else before moving on._

_Type your answer here, replacing this text._

<!-- END QUESTION -->

We are going to run the following hypothesis test to determine the association between winning and number of Star Baker awards. The population we are examining is every contestant from seasons 2 through 11 of GBBO. We are going to use the following null and alternative hypotheses:

**Null hypothesis:** The distribution of Star Baker awards between contestants who won their season and contestants who did not win their season is the same.

**Alternative hypothesis:** Contestants who win their season of the show will win more Star Baker awards on average.

Our alternative hypothesis is related to our suspicion that contestants who win more Star Baker awards are more skilled, so they are more likely to win the season.

The `bakers` table below describes the number of star baker awards each contest won and whether or not they won their season (`1` if they won, `0` if they did not win). The data was manually aggregated from Wikipedia for seasons 2-11 of the show. We randomized the order of rows as to not spoil the outcome of the show.

```python
bakers = Table.read_table("star_bakers.csv")
bakers.show(3)
```

## Task 02 📍

'star baker awards mean', 'won'Create a new table called `means` that contains the mean number of star baker awards for bakers who did not win (`won==0`) and bakers that did win (`won==1`). The table should have the column names `won` and `star baker awards mean`.

```python
means = ...
means
```

```python
grader.check("task_02")
```

## Task 03 📍🔎

<!-- BEGIN QUESTION -->

Visualize the distribution of Star Baker awards for winners and non-winners as overlaid histograms. Make sure to:

* Use the bins we provided.
* Use the group argument of `tbl.hist`. In order to produce several overlayed histograms based on unique values in a given column, we can do something like `tbl.hist(..., group=<col_name>, bins=...)`!

_We recommend that you check your histogram with someone before moving on._

```python
useful_bins = np.arange(0, 7)
...
```

<!-- END QUESTION -->

## Task 04 📍🔎

<!-- BEGIN QUESTION -->

We want to figure out if there is a difference between the distribution of Star Baker awards between winners and non winners. For the test statistic, subtract the mean star baker awards value for those that didn't win from the mean value for those that did win.

Which values of this test statistic support the alternative hypothesis?

_Before moving on, confirm your answer with someone._

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Task 05 📍

Set `observed_difference` to the observed test statistic using the `means` table.

```python
observed_difference = ...
observed_difference
```

```python
grader.check("task_05")
```

## Task 06 📍

Shortly, you will be calculating that difference on a variety of tables similar to `bakers`, but with the mean values shuffled around. Given a table like `bakers`, a label column `label_col` with two labels, `1` and `0`, and a values column `val_col`, write a function `find_test_stat` that calculates the test statistic outlined above.

```python
def find_test_stat(tbl, label_col, val_col):
    ...

# Confirm that your function produces the same observed difference
find_test_stat(bakers, "won", "star baker awards")
```

```python
grader.check("task_06")
```

When we run a simulation for A/B testing, we resample by **shuffling the labels** of the original sample. If the null hypothesis is true and the star baker award distributions are the same, we expect that the difference in mean star baker awards to not change when `"won"` labels are changed.

## Task 07 📍

Write a function `simulate_and_test_statistic` to compute one trial of our A/B test. Your function should run a simulation and return a test statistic.

```python
def simulate_and_test_statistic(tbl, labels_col, values_col):
    ...

simulate_and_test_statistic(bakers, "won", "star baker awards")
```

```python
grader.check("task_07")
```

## Task 08 📍

Simulate 5000 trials of our A/B test and store the test statistics in an array called `differences`. This test will take a few minutes to run.

```python
repetitions = ...
differences = ...

...                                              

differences
```

```python
grader.check("task_08")
```

Run the cell below to view a histogram of your simulated test statistics plotted with your observed test statistic.

```python
Table().with_column('Difference Between Group Means', differences).hist(bins=20)
plt.scatter(observed_difference, 0, color='red', s=30, zorder=2)
plt.ylim(-0.1, 1.35);
```

## Task 09 📍

Find the p-value for your test and assign it to `empirical_p`.

```python
empirical_p = ...                                             

empirical_p 
```

```python
grader.check("task_09")
```

## Task 10 📍🔎

<!-- BEGIN QUESTION -->

1. Using a 5% p-value cutoff, draw a conclusion about the null and alternative hypotheses using more the technical language of hypothesis testing.
2. Using simple, non-technical language, what does your analysis tell you about the association between star baker awards and winning? 
3. What can you claim about causation from your statistical analysis? 
    
Confirm your answer with a peer, instructor, etc.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Submit your Lab to Canvas

Once you have finished working on the lab questions, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the requirements for a Complete score for this lab assignments.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
4. Select the menu items "File", "Save and Export Notebook As...", and "HTML (.html)" in the notebook's Toolbar to download an HTML version of this notebook file.
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded HTML file.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_02 (points: None)
```
>>> 
>>> means.num_rows == 2
True
```
failure_message: ❌ Your table should have 2 rows.
success_message: ✅ Your table has 2 rows.
```
>>> 
>>> {'won', 'star baker awards mean', } == set(means.labels)
True
```
failure_message: ❌ Your labels should be 'won' and 'star baker awards mean'.
success_message: ✅ Your labels seem correct.
```
>>> 
>>> (np.round(min(means.column("star baker awards mean")), 2) == 0.65
...  and np.round(max(means.column("star baker awards mean")), 2) == 1.5)
True
```
failure_message: ❌ Your mean values don't seem correct.
success_message: ✅ Your mean values seem correct.


Tests for task_05 (points: None)
```
>>> 
>>> isinstance(observed_difference, float)
True
```
failure_message: ❌ Your observed_difference value should be a float.
success_message: ✅ Your observed_difference value is a float.
```
>>> 
>>> float(round(observed_difference, 3)) == 0.848
True
```
failure_message: ❌ Your observed_difference value does't seem correct. Double check that you are subtracing the means for the winners and losers. Also, double check the order of your subtraction.
success_message: ✅ Your observed_difference value seems correct.


Tests for task_06 (points: None)
```
>>> 
>>> np.isclose(round(find_test_stat(bakers, "won", "star baker awards"), 3) - 0.848, 0)
True
```
failure_message: ❌ Your function should produce the same observed value as in the previous task.
success_message: ✅ It seems like your function correctly produces the same observed value as in the previous task.
```
>>> shifted_indices = np.roll(np.arange(bakers.num_rows), 2)
>>> shifted_won_vals = bakers.take(shifted_indices).column('won')
>>> test_tbl = bakers.select('star baker awards').with_column('Shuffled Label', shifted_won_vals)
>>> np.isclose(round(find_test_stat(test_tbl, "Shuffled Label", "star baker awards"), 3) - (-0.132), 0)
True
```
failure_message: ❌ Your function should use the arguments tbl, label_col, val_col to do the calculation.
success_message: ✅ It seems like your function is using the arguments tbl, label_col, val_col to do the calculation.


Tests for task_07 (points: None)
```
>>> 
>>> test_stat = round(simulate_and_test_statistic(bakers, "won", "star baker awards"), 3)
>>> -2 < test_stat < 2
True
```
failure_message: ❌ Your test_stat value should be between -2 and 2. Make sure you are using find_test_stat.
success_message: ✅ Your test_stat value is between -2 and 2.
```
>>> 
>>> np.random.seed(1)
>>> test_stat2 = simulate_and_test_statistic(bakers, "won", "star baker awards")
>>> np.round(test_stat2, 3) == -0.023 or np.round(test_stat2, 3) == -0.132
True
```
failure_message: ❌ simulate_and_test_statistic doesn't seem to work as intended. Make sure you are using the sample method with the arguement with_replacement=False to suffle the labels. 
success_message: ✅ simulate_and_test_statistic seems to work for the table we checked.


Tests for task_08 (points: None)
```
>>> 
>>> len(differences)
5000
```
failure_message: ❌ differences should have 5000 items.
success_message: ✅ differences has 5000 items.
```
>>> 
>>> abs(np.average(differences)) < 0.05
True
```
failure_message: ❌ On average, your test statistic should be close to 0.
success_message: ✅ On average, your test statistic should is close to 0
```
>>> 
>>> all(differences == differences.item(0)) == False
True
```
failure_message: ❌ Make sure all of the test statistics are different
success_message: ✅ All of the test statistics are different


Tests for task_09 (points: None)
```
>>> 
>>> 0 <= empirical_p < 0.05
True
```
failure_message: ❌ Your p-value should be less than 5%. (It is rare, but there is a small chance you did everything correctly and you just ended up with a p-value that is 5% or larger. Re-run the code in Task 08 and Task 09 to double check.)
success_message: ✅ Your p-value should be less than 5%.
