```python
# Initialize Otter
import otter
grader = otter.Notebook("lab06.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

<div style="text-align: center;">
    <h1>Lab 06: Simulations</h1>
    <em>View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.</em>
</div>

This lab covers: [iteration](https://www.inferentialthinking.com/chapters/09/2/Iteration.html) and [simulations](https://www.inferentialthinking.com/chapters/09/3/Simulation.html), as well as introduce the concept of [randomness](https://www.inferentialthinking.com/chapters/09/Randomness.html).

The data used in this lab will contain salary data and other statistics for basketball players from the 2014-2015 NBA season. This data was collected from the following sports analytic sites: [Basketball Reference](http://www.basketball-reference.com) and [Spotrac](http://www.spotrac.com).

Run the following cell.

```python
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## Nachos and Conditionals

In Python, the boolean data type contains only two unique values:  `True` and `False`. Expressions containing comparison operators such as `<` (less than), `>` (greater than), and `==` (equal to) evaluate to Boolean values. A list of common comparison operators can be found below!

<img src="comparisons.png">

Run the cell below to see an example of a comparison operator in action.

```python
3 > 1 + 1
```

We can even assign the result of a comparison operation to a variable.

```python
result = 10 / 2 == 5
result
```

Arrays are compatible with comparison operators. The output is an array of boolean values.

```python
make_array(1, 5, 7, 8, 3, -1) > 3
```

One day, when you come home after a long week, you see a hot bowl of nachos waiting on the dining table! Let's say that whenever you take a nacho from the bowl, it will either have only **cheese**, only **salsa**, **both** cheese and salsa, or **neither** cheese nor salsa (a sad tortilla chip indeed). 

Let's try and simulate taking nachos from the bowl at random using the function, `np.random.choice(...)`.

### `np.random.choice`

`np.random.choice` picks one item at random from the given array. It is equally likely to pick any of the items. Run the cell below several times, and observe how the results change.

```python
nachos = make_array('cheese', 'salsa', 'both', 'neither')
np.random.choice(nachos)
```

To repeat this process multiple times, pass in an int `n` as the second argument to return `n` different random choices. By default, `np.random.choice` samples **with replacement** and returns an *array* of items. 

Run the next cell to see an example of sampling with replacement 10 times from the `nachos` array.

```python
np.random.choice(nachos, 10)
```

To count the number of times a certain type of nacho is randomly chosen, we can use `np.count_nonzero`

### `np.count_nonzero`

`np.count_nonzero` counts the number of non-zero values that appear in an array. When an array of boolean values are passed through the function, it will count the number of `True` values (remember that in Python, `True` is coded as 1 and `False` is coded as 0.)

Run the next cell to see an example that uses `np.count_nonzero`.

```python
np.count_nonzero(make_array(True, False, False, True, True))
```

### Task 01 📍

Assume we took ten nachos at random, and stored the results in an array called `ten_nachos` as done below. Find the number of nachos with only cheese using code (do not hardcode the answer).  

*Hint:* Our solution involves a comparison operator (e.g. `=`, `<`, ...) and the `np.count_nonzero` method.


```python
ten_nachos = make_array('neither', 'cheese', 'both', 'both', 'cheese', 'salsa', 'both', 'neither', 'cheese', 'both')
number_cheese = ...
number_cheese
```

```python
grader.check("task_01")
```

### Conditional Statements

A conditional statement is a multi-line statement that allows Python to choose among different alternatives based on the truth value of an expression.

Here is a basic example.

```
def sign(x):
    if x > 0:
        return 'Positive'
    else:
        return 'Negative'
```

If the input `x` is greater than `0`, we return the string `'Positive'`. Otherwise, we return `'Negative'`.

If we want to test multiple conditions at once, we use the following general format.

```
if <if expression>:
    <if body>
elif <elif expression 0>:
    <elif body 0>
elif <elif expression 1>:
    <elif body 1>
...
else:
    <else body>
```

Only the body for the first conditional expression that is true will be evaluated. Each `if` and `elif` expression is evaluated and considered in order, starting at the top. As soon as a true value is found, the corresponding body is executed, and the rest of the conditional statement is skipped. If none of the `if` or `elif` expressions are true, then the `else body` is executed. 

For more examples and explanation, refer to the section on conditional statements [here](https://inferentialthinking.com/chapters/09/1/Conditional_Statements.html).

### Task 02 📍

Complete the following conditional statement so that the string `'More please'` is assigned to the variable `say_please` if the number of nachos with cheese in `ten_nachos` is less than `5`.

*Hint*: You should be using `number_cheese` from Question 1.


```python
say_please = '?'

if ...:
    say_please = 'More please' 

say_please
```

```python
grader.check("task_02")
```

### Task 03 📍

Write a function called `nacho_reaction` that returns a reaction (as a string) based on the type of nacho passed in as an argument. Use the table below to match the nacho type to the appropriate reaction.

<img src="nacho_reactions.png">

*Hint:* If you're failing the test, double check the spelling of your reactions.


```python
def nacho_reaction(nacho):
    if nacho == "cheese":
        ...
    ... :
        ...
    ... :
        ...
    ... :
        ...


spicy_nacho = nacho_reaction('salsa')
spicy_nacho
```

```python
grader.check("task_03")
```

### Task 04 📍

Create a table `ten_nachos_reactions` that consists of the nachos in `ten_nachos` as well as the reactions for each of those nachos. The columns should be called `Nachos` and `Reactions`.

*Hint:* Use the `apply` method. 


```python
ten_nachos_tbl = Table().with_column('Nachos', ten_nachos)
...
ten_nachos_reactions
```

```python
grader.check("task_04")
```

### Task 05 📍

Using code, find the number of 'Wow!' reactions for the nachos in `ten_nachos_reactions`.


```python
number_wow_reactions = ...
number_wow_reactions
```

```python
grader.check("task_05")
```

## Simulations and For Loops

Using a `for` statement, we can perform a task multiple times. This is known as iteration.

One use of iteration is to loop through a set of values. For instance, we can print out all of the colors of the rainbow.

```python
rainbow = make_array("red", "orange", "yellow", "green", "blue", "indigo", "violet")

for color in rainbow:
    print(color)
```

We can see that the indented part of the `for` loop, known as the body, is executed once for each item in `rainbow`. The name `color` is assigned to the next value in `rainbow` at the start of each iteration. Note that the name `color` is arbitrary; we could easily have named it something else. The important thing is we stay consistent throughout the `for` loop. 

```python
for another_name in rainbow:
    print(another_name)
```

In general, however, we would like the variable name to be somewhat informative. 

### Task 06 📍

In the following cell, we've loaded the text of _Pride and Prejudice_ by Jane Austen, split it into individual words, and stored these words in an array `p_and_p_words`. Using a `for` loop, assign `longer_than_five` to the number of words in the novel that are more than 5 letters long.

*Hint*: You can find the number of letters in a word with the `len` function.


```python
austen_string = open('Austen_PrideAndPrejudice.txt', encoding='utf-8').read()
p_and_p_words = np.array(austen_string.split())

longer_than_five = ...

# a for loop would be useful here



longer_than_five
```

```python
grader.check("task_06")
```

### Task 07 📍

Using a simulation with 10,000 trials, assign num_different to the number of times, in 10,000 trials, that two words picked uniformly at random (with replacement) from Pride and Prejudice have different lengths. 

*Hint 1*: What function did we use in section 1 to sample at random with replacement from an array? 

*Hint 2*: Remember that `!=` checks for non-equality between two items.


```python
trials = 10000
num_different = ...

for ... in ...:
    ...

num_different
```

```python
grader.check("task_07")
```

We can also use `np.random.choice` to simulate multiple trials.

### Task 08 📍

Allie is playing darts. Her dartboard contains ten equal-sized zones with point values from 1 to 10. Write code that simulates her total score after 1000 dart tosses.

*Hint:* First decide the possible values you can take in the experiment (point values in this case). Then use `np.random.choice` to simulate Allie's tosses. Finally, sum up the scores to get Allie's total score.


```python
num_tosses = 1000
simulated_tosses = ...
total_score = ...
total_score
```

```python
grader.check("task_08")
```

## Sampling Basketball Data

We will now introduce the topic of sampling, which we’ll be discussing in more depth in this week’s lectures. We’ll guide you through this code, but if you wish to read more about different kinds of samples before attempting this question, you can check out [section 10 of the textbook](https://www.inferentialthinking.com/chapters/10/Sampling_and_Empirical_Distributions.html).

Run the cell below to load player and salary data that we will use for our sampling. 

```python
player_data = Table().read_table("player_data.csv")
salary_data = Table().read_table("salary_data.csv")
full_data = salary_data.join("PlayerName", player_data, "Name")

# The show method immediately displays the contents of a table. 
# This way, we can display the top of two tables using a single cell.
player_data.show(3)
salary_data.show(3)
full_data.show(3)
```

Rather than getting data on every player (as in the tables loaded above), imagine that we had gotten data on only a smaller subset of the players. For 492 players, it's not so unreasonable to expect to see all the data, but usually we aren't so lucky. 

If we want to make estimates about a certain numerical property of the population (known as a statistic, e.g. the mean or median), we may have to come up with these estimates based only on a smaller sample. Whether these estimates are useful or not often depends on how the sample was gathered. We have prepared some example sample datasets to see how they compare to the full NBA dataset. Later we'll ask you to create your own samples to see how they behave.

To save typing and increase the clarity of your code, we will package the analysis code into a few functions. This will be useful in the rest of the lab as we will repeatedly need to create histograms and collect summary statistics from that data.

We've defined the `histograms` function below, which takes a table with columns `Age` and `Salary` and draws a histogram for each one. It uses bin widths of 1 year for `Age` and $1,000,000 for `Salary`.

```python
def histograms(t):
    ages = t.column('Age')
    salaries = t.column('Salary')/1000000
    t1 = t.drop('Salary').with_column('Salary', salaries)
    age_bins = np.arange(min(ages), max(ages) + 2, 1) 
    salary_bins = np.arange(min(salaries), max(salaries) + 1, 1)
    t1.hist('Age', bins=age_bins, unit='year')
    plt.title('Age distribution')
    t1.hist('Salary', bins=salary_bins, unit='million dollars')
    plt.title('Salary distribution') 
    
histograms(full_data)
print('Two histograms should be displayed below')
```

### Task 09 📍

Create a function called `compute_statistics` that takes a table containing ages and salaries and:
- Draws a histogram of ages
- Draws a histogram of salaries
- Returns a two-element array containing the average age and average salary (in that order)

You can call the `histograms` function to draw the histograms! 

*Note:* More charts will be displayed when running the test cell. Please feel free to ignore the charts.


```python
def compute_statistics(age_and_salary_data):
    ...
    age = ...
    salary = ...
    ...
    

full_stats = compute_statistics(full_data)
full_stats
```

```python
grader.check("task_09")
```

### Simple random sampling

A more justifiable approach is to sample uniformly at random from the players.  In a **simple random sample (SRS) without replacement**, we ensure that each player is selected at most once. Imagine writing down each player's name on a card, putting the cards in an box, and shuffling the box.  Then, pull out cards one by one and set them aside, stopping when the specified sample size is reached.

### Producing simple random samples

Sometimes, it’s useful to take random samples even when we have the data for the whole population. It helps us understand sampling accuracy.

### `sample`

The table method `sample` produces a random sample from the table. By default, it draws at random **with replacement** from the rows of a table. It takes in the sample size as its argument and returns a **table** with only the rows that were selected. 

Run the cell below to see an example call to `sample()` with a sample size of 5, with replacement.

```python
# Just run this cell

salary_data.sample(5)
```

The optional argument `with_replacement=False` can be passed through `sample()` to specify that the sample should be drawn without replacement.

Run the cell below to see an example call to `sample()` with a sample size of 5, without replacement.

```python
# Just run this cell

salary_data.sample(5, with_replacement=False)
```

### Task 10 📍🔎

<!-- BEGIN QUESTION -->

Produce a simple random sample of size 44 from `full_data`. Run your analysis on it again.  Run the cell a few times to see how the histograms and statistics change across different samples.

- How much does the average age change across samples? 
- What about average salary?


_Type your answer here, replacing this text._

```python
my_small_srswor_data = ...
my_small_stats = ...
my_small_stats
```

<!-- END QUESTION -->

### Task 11 📍🔎

As in the previous question, analyze several simple random samples of size 100 from `full_data`.  
- Do the histogram shapes seem to change more or less across samples of 100 than across samples of size 44?  
- Are the sample averages and histograms closer to their true values/shape for age or for salary?  What did you expect to see?


_Type your answer here, replacing this text._

```python
my_large_srswor_data = ...
my_large_stats = ...
my_large_stats
```

## Submit your Lab to Canvas

Once you have finished working on the lab questions, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the requirements for a Complete score for this lab assignments.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
4. Select the menu items "File", "Save and Export Notebook As...", and "HTML (.html)" in the notebook's Toolbar to download an HTML version of this notebook file.
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded HTML file.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: None)
```
>>> 
>>> number_cheese == 3
True
```
failure_message: ❌ There are 3 nachos with only cheese. Try using np.count_nonzero.
success_message: ✅ There are 3 nachos with only cheese.


Tests for task_02 (points: None)
```
>>> 
>>> say_please == 'More please'
True
```
failure_message: ❌ Since number_cheese < 5, say_please should be assigned to 'More please'. Check if number_cheese < 5.
success_message: ✅ say_please is assigned to 'More please'.


Tests for task_03 (points: None)
```
>>> 
>>> nacho_reaction('salsa') == 'Spicy!'
True
```
failure_message: ❌ nacho_reaction('salsa') should display 'Spicy!' You might need to check your spelling/punctuation.
success_message: ✅ nacho_reaction('salsa') should display 'Spicy!'.
```
>>> 
>>> nacho_reaction('cheese') == 'Cheesy!'
True
```
failure_message: ❌ nacho_reaction('cheese') should display 'Cheesy!' You might need to check your spelling/punctuation.
success_message: ✅ nacho_reaction('cheese') should display 'Cheesy!'.
```
>>> 
>>> nacho_reaction('both') == 'Wow!'
True
```
failure_message: ❌ nacho_reaction('both') should display 'Wow!' You might need to check your spelling/punctuation.
success_message: ✅ nacho_reaction('both') should display 'Wow!'.
```
>>> 
>>> nacho_reaction('neither') == 'Meh.'
True
```
failure_message: ❌ nacho_reaction('neither') should display 'Meh.' You might need to check your spelling/punctuation.
success_message: ✅ nacho_reaction('neither') should display 'Meh.'.


Tests for task_04 (points: None)
```
>>> 
>>> # One or more of the reaction results could be incorrect;
>>> np.count_nonzero(ten_nachos_reactions.column('Reactions') == make_array('Meh.', 'Cheesy!', 'Wow!', 'Wow!', 'Cheesy!', 'Spicy!', 'Wow!', 'Meh.', 'Cheesy!', 'Wow!')) == 10
True
```
failure_message: ❌ It seems like one or more of the reaction results could be incorrect. To get the values for the table, trying using the apply method with the function you just made.
success_message: ✅ The reaction values in the table seem correct.


Tests for task_05 (points: None)
```
>>> 
>>> 2 < number_wow_reactions < 6
True
```
failure_message: ❌ number_wow_reactions should be between 2 and 6, non-inclusive. Try using np.count_nonzero.
success_message: ✅ number_wow_reactions is between 2 and 6, non-inclusive.
```
>>> 
>>> # Incorrect value for number_wow_reactions
>>> number_wow_reactions == 4
True
```
failure_message: ❌ You have the incorrect value for number_wow_reactions.
success_message: ✅ You have the correct value for number_wow_reactions.


Tests for task_06 (points: None)
```
>>> 
>>> longer_than_five == 35453
True
```
failure_message: ❌ The value assigned to longer_than_five is not correct. Make sure to initilalize longer_than_five to 0 and check if the length of the word is great than 5 as you loop through each word in p_and_p_words.
success_message: ✅ The value assigned to longer_than_five is correct.


Tests for task_07 (points: None)
```
>>> 
>>> 8100 <= num_different <= 9100
True
```
failure_message: ❌ num_different should be between 8100 and 9100. Try initializing num_different to 0, iterating through the number of trials, randomly choosing 2 words from p_and_p_words, checking the length of the 2 words, and updating num_different depending on the results of that check.
success_message: ✅ num_different represents a value betwen 8100 and 9100.


Tests for task_08 (points: None)
```
>>> 
>>> 1000 <= total_score <= 10000
True
```
failure_message: ❌ total_score should represent a value betwen 1000 and 10000. You might try using np.random.choice with an array of possible values and a sample size of num_tosses. The total_Score should be found by summing up all the values in simulated_tosses.
success_message: ✅ total_score represents a value betwen 1000 and 10000.


Tests for task_09 (points: None)
```
>>> 
>>> stats = compute_statistics(full_data)
>>> plt.close()
>>> plt.close()
>>> round(float(stats[0]), 2) == 26.54
True
```
failure_message: ❌ The age value we checked is not correct. Make sure you are using np.mean with the Age column data.
success_message: ✅ The age value we checked is correct.
```
>>> 
>>> stats = compute_statistics(full_data)
>>> plt.close()
>>> plt.close()
>>> round(float(stats[1]), 2) == 4269775.77
True
```
failure_message: ❌ The age value we checked is not correct. Make sure you are using np.mean with the Salary column data.
success_message: ✅ The salary value we checked is correct.
