```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

<div style="text-align: center;">
    <h1>Lab 04: Functions</h1>
    <em>View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.</em>
</div>

This week, we'll learn about functions and table methods such as `apply`! 

Recommended Reading:[Applying a Function to a Column](https://inferentialthinking.com/chapters/08/1/Applying_a_Function_to_a_Column.html)

First, set up the notebook by running the cell below.

```python
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')

from ipywidgets import interact, interactive, fixed, interact_manual
import ipywidgets as widgets
```

## Defining functions

Let's write a very simple function that converts a proportion to a percentage by multiplying it by 100.  For example, the value of `to_percentage(.5)` should be the number 50 (no percent sign).

A function definition has a few parts.

### `def`

It always starts with `def` (short for **def**ine):

``` python
def
```

### Name

Next comes the name of the function.  Like other names we've defined, it can't start with a number or contain spaces. Let's call our function `to_percentage`:

``` python
def to_percentage
```

### Signature

Next comes something called the *signature* of the function.  This tells Python how many arguments your function should have, and what names you'll use to refer to those arguments in the function's code.  A function can have any number of arguments (including 0!). 

`to_percentage` should take one argument, and we'll call that argument `proportion` since it should be a proportion.

``` python
def to_percentage(proportion)
```

If we want our function to take more than one argument, we add a comma between each argument name. Note that if we had zero arguments, we'd still place the parentheses () after than name. 

We put a colon after the signature to tell Python it's over. If you're getting a syntax error after defining a function, check to make sure you remembered the colon!

``` python
def to_percentage(proportion):
```

### Documentation

Functions can do complicated things, so you should write an explanation of what your function does.  For small functions, this is less important, but it's a good habit to learn from the start.  Conventionally, Python functions are documented by writing an **indented** triple-quoted string:

``` python
def to_percentage(proportion):
    """Converts a proportion to a percentage."""
```

### Body

Now we start writing code that runs when the function is called.  This is called the *body* of the function and every line **must be indented with a tab**.  Any lines that are *not* indented and left-aligned with the def statement is considered outside the function. 

Some notes about the body of the function:
- We can write code that we would write anywhere else.  
- We use the arguments defined in the function signature. We can do this because we assume that when we call the function, values are already assigned to those arguments.
- We generally avoid referencing variables defined *outside* the function. If you would like to reference variables outside of the function, pass them through as arguments!


Now, let's give a name to the number we multiply a proportion by to get a percentage:

``` python
def to_percentage(proportion):
    """Converts a proportion to a percentage."""
    factor = 100
```

### `return`

The special instruction `return` is part of the function's body and tells Python to make the value of the function call equal to whatever comes right after `return`.  We want the value of `to_percentage(.5)` to be the proportion .5 times the factor 100, so we write:

``` python
def to_percentage(proportion):
    """Converts a proportion to a percentage."""
    factor = 100
    return proportion * factor
```
        
`return` only makes sense in the context of a function, and **can never be used outside of a function**. `return` is always the last line of the function because Python stops executing the body of a function once it hits a `return` statement.

*Note:*  `return` inside a function tells Python what value the function evaluates to. However, there are other functions, like `print`, that have no `return` value. For example, `print` simply prints a certain value out to the console. 

`return` and `print` are **very** different. 

### Task 01 📍

1. Define `to_percentage` in the cell below.
2. Call your function to convert the proportion 0.2 to a percentage and name the output `twenty_percent`.


```python
def ...
    """"
    Converts a proportion to a percentage.
    
    >>> to_percentage(0.25)
    25.0
    """
    return ...
    
twenty_percent = ...
twenty_percent

```

```python
grader.check("task_01")
```

Like you’ve done with built-in functions in previous labs (max, abs, etc.), you can pass in named values as arguments to your function.

### Task 02 📍

Use `to_percentage` again to convert the proportion named `a_proportion` (defined below) to a percentage called `a_percentage`.

*Note:* You don't need to define `to_percentage` again!  Like other named values, functions stick around after you define them.


```python
a_proportion = 2 ** (0.5) / 2
a_percentage = ...
a_percentage
```

```python
grader.check("task_02")
```

Here's something important about functions: the names assigned *within* a function body are only accessible within the function body. Once the function has returned, those names are gone.  So even if you created a variable called `factor` and defined `factor = 100` inside of the body of the `to_percentage` function and then called `to_percentage`, `factor` would not have a value assigned to it outside of the body of `to_percentage`.

As we've seen with built-in functions, functions can also take strings (or arrays, or tables) as arguments, and they can return those things, too.

### Task 03 📍

Define a function called `disemvowel`.  It should take a single string as its argument.  (You can call that argument whatever you want.)  It should return a copy of that string, but with all the characters that are vowels removed.  (In English, the vowels are the characters "a", "e", "i", "o", and "u".) You can use as many lines inside of the function to do this as you’d like.

*Hint:* To remove all the "a"s from a string, you can use `that_string.replace("a", "")`.  The `.replace` method for strings returns a new string, so you can call `replace` multiple times, one after the other. 


```python
def disemvowel(a_string):
    """
    Removes all vowels from a string.
    
    >>> disemvowel('datascience')
    'dtscnc'
    """ 
    ...

# Checking to see if your function works for an example string.
disemvowel("Can you read this without vowels?")
```

```python
grader.check("task_03")
```

### Calls on calls on calls

Just as you write a series of lines to build up a complex computation, it's useful to define a series of small functions that build on each other.  Since you can write any code inside a function's body, you can call other functions you've written.

If a function is a like a recipe, defining a function in terms of other functions is like having a recipe for cake telling you to follow another recipe to make the frosting, and another to make the jam filling.  This makes the cake recipe shorter and clearer, and it avoids having a bunch of duplicated frosting recipes.  It's a foundation of productive programming.

For example, suppose you want to count the number of characters *that aren't vowels* in a piece of text.  One way to do that is this to remove all the vowels and count the size of the remaining string.

### Task 04 📍

Write a function called `num_non_vowels`.  It should take a string as its argument and return a number.  That number should be the number of characters in the argument string that aren't vowels. You should use the `disemvowel` function you wrote above inside of the `num_non_vowels` function.

*Hint:* The function `len` takes a string as its argument and returns the number of characters in it.


```python
def num_non_vowels(a_string):
    """
    The number of characters in a string, minus the vowels.
    
    >>> num_non_vowels('datascience')
    
    6
    """
    ...

# Try running an example function call below to test out your code.

```

```python
grader.check("task_04")
```

Functions can also encapsulate code that *displays output* instead of computing a value. For example, if you call `print` inside a function, and then call that function, something will get printed.

The `top_movies_1995_2022.csv` dataset has information about movie sales in recent years.  Suppose you'd like to display the year with the 5th-highest total gross (adjusted) movie sales, printed in a human-readable way.  You might do this:

```python
movies_by_year = Table.read_table("top_movies_1995_2022.csv")
k = 5
fifth_from_top_movie_year = movies_by_year.sort("Total Gross (Adjusted)", descending=True).column("Year").item(k-1)
print("Year number", k, "for total gross movie sales was:", fifth_from_top_movie_year)
```

After writing this, you realize you also wanted to print out the 2nd and 3rd-highest years.  Instead of copying your code, you decide to put it in a function.  Since the rank varies, you make that an argument to your function.

### Task 05 📍

Write a function called `print_kth_top_movie_year`.

* It should take a single argument, `k`, the rank of the year (like 2, 3, or 5 in the above examples).  
* It should print out a message like the one above.  So, rather than a `return` statement, your function will have a `print` statement at the end.

```python
...
    print(...)

...

# Try running an example function call below with k = 2 to test out your code.
print_kth_top_movie_year(2)
```

```python
grader.check("task_05")
```

`interact` from the `ipywidgets` module allows you to pass in an array for a function argument. It will then present a dropdown menu of options. Run the following cell to see how it works.

`interact` produces an output that isn't helpful in this context. Some people use a symbol like the underscore `_` for the name of something that produces an unwanted output to prevent form showing.

```python
_ = interact(print_kth_top_movie_year, k=np.arange(1, 10))
```

### `print` is not the same as `return`

The `print_kth_top_movie_year(k)` function prints the total gross movie sales for the year that was provided! However, since we did not return any value in this function, we can not use it after we call it. Let's look at an example of another function that prints a value but does not return it.

```python
def print_number_five():
    print(5)
```

```python
print_number_five()
```

However, if we try to use the output of `print_number_five()`, we see that the value `5` is printed but we get a TypeError when we try to add the number 2 to it!

<img src="./nonetype_error.png" alt="The TypeError seen when using the named output of print statement.">

It may seem that `print_number_five()` is returning a value, 5. In reality, it just displays the number 5 to you without giving you the actual value! If your function prints out a value without returning it and you try to use that value, you will run into errors, so be careful!

Explain to your neighbor how you might add a line of code to the `print_number_five` function (after `print(5)`) so that the code `print_number_five_output + 5` would result in the value `10`, rather than an error.

## Functions and CEO Incomes

In this question, we'll look at the 2015 compensation of CEOs at the 100 largest companies in California. The data was compiled from a [Los Angeles Times analysis](http://spreadsheets.latimes.com/california-ceo-compensation/), and ultimately came from [filings](https://www.sec.gov/answers/proxyhtf.htm) mandated by the SEC from all publicly-traded companies. Two companies have two CEOs, so there are 102 CEOs in the dataset.

We've copied the raw data from the LA Times page into a file called `raw_compensation.csv`. (The page notes that all dollar amounts are in **millions of dollars**.)

```python
raw_compensation = Table.read_table('raw_compensation.csv')
raw_compensation
```

We want to compute the average of the CEOs' pay. View the error produced from the following code:

<img src="./totalpay_error.png" alt="The UFuncTypeError the occurs from running np.average(raw_compensation.column('Total Pay'))">

Let's examine why this error occurred by looking at the values in the `Total Pay` column. 

### Task 06 📍

Use the `type` function and set `total_pay_type` to the type of the first value in the "Total Pay" column.


```python
total_pay_type = ...
total_pay_type
```

```python
grader.check("task_06")
```

### Task 07 📍

You should have found that the values in the `Total Pay` column are strings. It doesn't make sense to take the average of string values, so we need to convert them to numbers if we want to do this. Extract the first value in `Total Pay`.  It's Mark Hurd's pay in 2015, in *millions* of dollars.  Call it `mark_hurd_pay_string`.


```python
mark_hurd_pay_string = ...
mark_hurd_pay_string
```

```python
grader.check("task_07")
```

### Task 08 📍

Convert `mark_hurd_pay_string` to a number of *dollars*. 

Some hints, as this question requires multiple steps:
- The string method `strip` will be useful for removing the dollar sign; it removes a specified character from the start or end of a string.  For example, the value of `"100%".strip("%")` is the string `"100"`.  
- You'll also need the function `float`, which converts a string that looks like a number to an actual number.  
- Finally, remember that the answer should be in dollars, not millions of dollars.


```python
mark_hurd_pay = ...
mark_hurd_pay
```

```python
grader.check("task_08")
```

To compute the average pay, we need to do this for every CEO.  But that looks like it would involve copying this code 102 times.

This is where functions come in.  First, we'll define a new function, giving a name to the expression that converts "total pay" strings to numeric values.  Later in this lab, we'll see the payoff: we can call that function on every pay string in the dataset at once.

The next section of this lab explains how to define a function For now, just fill in the ellipses in the cell below.

### Task 09 📍

Copy the expression you used to compute `mark_hurd_pay`, and use it as the return expression of the function below. But make sure you replace the specific `mark_hurd_pay_string` with the generic `pay_string` name specified in the first line in the `def` statement.

*Hint*: When dealing with functions, you should generally not be referencing any variable outside of the function. Usually, you want to be working with the arguments that are passed into it, such as `pay_string` for this function. If you're using `mark_hurd_pay_string` within your function, you're referencing an outside variable!


```python
def convert_pay_string_to_number(pay_string):
    """
    Converts a pay string like '$100' (in millions) to a number of
    dollars.
    
    >>> convert_pay_string_to_number("$100 ")
    100000000.0
    """
    ...
```

```python
grader.check("task_09")
```

Running that cell doesn't convert any particular pay string. Instead, it creates a function called `convert_pay_string_to_number` that can convert *any* string with the right format to a number representing millions of dollars.

We can call our function just like we call the built-in functions we've seen. It takes one argument -- a string -- and it returns a float.

```python
convert_pay_string_to_number('$42')
```

```python
convert_pay_string_to_number(mark_hurd_pay_string)
```

We can also compute Safra Catz's pay in the same way. Run the following cell to see the results.

```python
convert_pay_string_to_number(raw_compensation.where("Name", are.containing("Safra")).column("Total Pay").item(0))
```

So, what have we gained by defining the `convert_pay_string_to_number` function? 
Well, without it, we'd have to copy the code `10**6 * float(some_pay_string.strip("$"))` each time we wanted to convert a pay string.  Now we just call a function whose name says exactly what it's doing.

## `apply`ing functions

Defining a function is a lot like giving a name to a value with `=`.  In fact, a function is a value just like the number 1 or the text "data"!

For example, we can make a new name for the built-in function `max` if we want:

```python
our_name_for_max = max
our_name_for_max(2, 6)
```

The old name for `max` is still around:

```python
max(2, 6)
```

Try just writing `max` or `our_name_for_max` (or the name of any other function) in a cell, and run that cell.  Python will print out a (very brief) description of the function.

```python
max
```

Now try writing `?max` or `?our_name_for_max` (or the name of any other function) in a cell, and run that cell.  A information box should show up at the bottom of your screen a longer description of the function

*Note: You can also press Shift+Tab after clicking on a name to see similar information!*

```python
?our_name_for_max
```

Let's look at what happens when we set `max`to a non-function value. You'll notice that a TypeError will occur when you try calling `max`. Things like integers and strings are not callable. Look out for any functions that might have been renamed when you encounter this type of error

<img src="./rename_function_error.png" alt="The TypeError that occurs when max is used to name a value and then used as a function.">

The following cell resets `max` to the built-in function. Just run this cell to reset `max` just in case you changed it.

```python
#Just run this cell, don't change its contents
import builtins
max = builtins.max
```

Why is this useful?  Since functions are just values, it's possible to pass them as arguments to other functions.  Here's a simple but not-so-practical example: we can make an array of functions.

```python
make_array(max, np.average, are.equal_to)
```

### Task 10 📍

Make an array containing any 3 other functions you've seen.  Call it `some_functions`.


```python
some_functions = ...
some_functions
```

```python
grader.check("task_10")
```

Working with functions as values can lead to some funny-looking code. For example, see if you can figure out why the following code works. Check your explanation with a neighbor or a staff member.

```python
make_array(max, np.average, are.equal_to).item(0)(4, -2, 7)
```

A more useful example of passing functions to other functions as arguments is the table method `apply`.

`apply` calls a function many times, once on *each* element in a column of a table.  It produces an *array* of the results.  Here we use `apply` to convert every CEO's pay to a number, using the function you defined:

```python
raw_compensation.apply(convert_pay_string_to_number, "Total Pay")
```

Here's an illustration of what that did:

<img src="apply.png"/>

Note that we didn’t write `raw_compensation.apply(convert_pay_string_to_number(), “Total Pay”)` or `raw_compensation.apply(convert_pay_string_to_number(“Total Pay”))`. We just passed the name of the function, with no parentheses, to `apply`, because all we want to do is let `apply` know the name of the function we’d like to use and the name of the column we’d like to use it on. `apply` will then call the function `convert_pay_string_to_number` on each value in the column for us!

### Task 11 📍

Using `apply`, make a table that's a copy of `raw_compensation` with one additional column called `Total Pay ($)`.  That column should contain the result of applying `convert_pay_string_to_number` to the `Total Pay` column (as we did above).  Call the new table `compensation`.


```python
compensation = raw_compensation.with_column(
    "Total Pay ($)",
    ...
    ) 
compensation
```

```python
grader.check("task_11")
```

Now that we have all the pays as numbers, we can learn more about them through computation.

### Task 12 📍

Compute the average total pay of the CEOs in the dataset.


```python
average_total_pay = ...
average_total_pay
```

```python
grader.check("task_12")
```

### Task 13 📍

Companies pay executives in a variety of ways: in cash, by granting stock or other equity in the company, or with ancillary benefits (like private jets).  Compute the proportion of each CEO's pay that was cash.  (Your answer should be an array of numbers, one for each CEO in the dataset.)

*Note:* When you answer this question, you'll encounter a red box appearing below your code cell that says something like `RuntimeWarning: invalid value encountered in true_divide`. Don't worry too much about the message. Warnings are raised by Python when it encounters an unusual condition in your code, but the condition is not severe enough to warrant throwing an error. 

The warning below is Python's cryptic way of telling you that you're dividing a number by zero. If you extract the values in `Total Pay ($)` as an array, you'll see that the last element is 0.


```python
cash_proportion = ...
cash_proportion
```

```python
grader.check("task_13")
```

Check out the `% Change` column in `compensation`.  It shows the percentage increase in the CEO's pay from the previous year.  For CEOs with no previous year on record, it instead says "(No previous year)".  The values in this column are *strings*, not numbers, so like the `Total Pay` column, it's not usable without a bit of extra work.

Given your current pay and the percentage increase from the previous year, you can compute your previous year's pay.  For example, if your pay is $\$120$ this year, and that's an increase of 50% from the previous year, then your previous year's pay was $\frac{\$120}{1 + \frac{50}{100}}$, or \$80.

### Task 14 📍

Create a new table called `with_previous_compensation`.  It should be a copy of `compensation`, but with the "(No previous year)" CEOs filtered out, and with an extra column called `2014 Total Pay ($)`.  That column should have each CEO's pay in 2014.

We've organized the template for this answer to guide you in roughly 5 steps. (If our steps are confusing you, you can delete the current structure and approach the problem your own way!)

1. Define a function percent_string_to_num to turn a percent to a number.
2. Create a compensation table having_previous_year where there is a previous year.
3. Get the percent changes, named percent_changes, as numbers instead of strings. We're still working off the table having_previous_year
4. Calculate the previous year's pay and name it previous_pay. We're still working off the table having_previous_year
5. Put the previous pay column into the having_previous_year table.

I can be helpful to print out your results after each step to make sure you're on the right track.

```python
# Step 1
def percent_string_to_num(percent_string):
    """Converts a percentage string to a number."""
    return ...

# Step 2
having_previous_year = ...

# Step 3
percent_changes = ...

# Step 4
previous_pay = ...

# Step 5
with_previous_compensation = ...

```

```python
grader.check("task_14")
```

### Task 15 📍

What was the average pay of these CEOs in 2014?


```python
average_pay_2014 = ...
average_pay_2014
```

```python
grader.check("task_15")
```

### Why is `apply` useful?

For operations like arithmetic, or the functions in the NumPy library, you don't need to use `apply`, because they automatically work on each element of an array.  But there are many things that don't.  The string manipulation we did in today's lab is one example.  Since you can write any code you want in a function, `apply` gives you total control over how you operate on data.

## Submit your Lab to Canvas

Once you have finished working on the lab questions, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the requirements for a Complete score for this lab assignments.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
4. Select the menu items "File", "Save and Export Notebook As...", and "HTML (.html)" in the notebook's Toolbar to download an HTML version of this notebook file.
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded HTML file.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: None)
```
>>> 
>>> to_percentage(0.25) == 25.0
True
```
failure_message: ❌ Make sure you are mulitplying the input proportion 100.
success_message: ✅ It seems like your function is working!


Tests for task_02 (points: None)
```
>>> 
>>> a_percentage == 70.71067811865476
True
```
failure_message: ❌ It doesn't seem like you are calling the function to_percentage on a_proportion.
success_message: ✅ Great work calling the function on the value named a_proportion.


Tests for task_03 (points: None)
```
>>> 
>>> disemvowel("Datascience rules!") == "Dtscnc rls!"
True
```
failure_message: ❌ It doesn't seem like disemvowel is wokring correctly. a_string.replace('t', '') will remove any t's in the string named a_string.
success_message: ✅ It looks like disemvowel is working correctly.


Tests for task_04 (points: None)
```
>>> 
>>> num_non_vowels("datascience") == 6
True
```
failure_message: ❌ It doesn't seem like your function num_non_vowels is working correctly. Try using len on the output of disemvowel.
success_message: ✅ It seems like your function num_non_vowels is working correctly!


Tests for task_05 (points: None)
```
>>> 
>>> print(print_kth_top_movie_year(2))
Year number 2 for total gross movie sales was: 2019
None
```
failure_message: ❌ It doesn't seem like your function print_kth_top_movie_year is working correctly. Try using the sort and column table methods along with the array method item. Don't forget the index number is 1 less than the item number in an array
success_message: ✅ It seems like your function print_kth_top_movie_year is working correctly!


Tests for task_06 (points: None)
```
>>> 
>>> import numpy
>>> total_pay_type != numpy.ndarray
True
```
failure_message: ❌ Make sure you are examining the values in the column, not the column itself.
success_message: ✅ Great work naming the item and not the array.
```
>>> 
>>> 'str' in str(total_pay_type)
True
```
failure_message: ❌ You should get a string if you get the type of the frist item in the Total Pay column.
success_message: ✅ Nice! The type of the frist item in the Total Pay column is a string.
```
>>> 
>>> total_pay_type != int
True
```
failure_message: ❌ The value might look like a number, but it isn't. Make sure to call the type function on a value in the column.
success_message: ✅ Looks good! The value might look like a number, but it isn't.


Tests for task_07 (points: None)
```
>>> 
>>> str(mark_hurd_pay_string) == '$53.25 '
True
```
failure_message: ❌ It doesn't seem like you have the correct value. You want to get the first item in the Total Pay column's array of values.
success_message: ✅ Nice work getting the correct value.


Tests for task_08 (points: None)
```
>>> 
>>> # Your answer should be a number
>>> type(mark_hurd_pay) != str
True
```
failure_message: ❌ The value is still a string. Try stripping the $ sign and using float.
success_message: ✅ It looks like the value is no longer a string.
```
>>> 
>>> mark_hurd_pay != 5325
True
```
failure_message: ❌ Don't forget to give your answer in dollars, not millions of dollars. Try multiplying by one million.
success_message: ✅ Nice! You don't have the dolloar amoutn in millions of dollars.
```
>>> 
>>> mark_hurd_pay == 53250000
True
```
failure_message: ❌ Your value doesn't seem correct.
success_message: ✅ It looks like you have the correct value.


Tests for task_09 (points: None)
```
>>> 
>>> convert_pay_string_to_number("$100 ") == 100000000.0
True
```
failure_message: ❌ convert_pay_string_to_number is not working for our example. Try to redo the strips from a previous task, but use the name pay_string for the string.
success_message: ✅ It looks like your function convert_pay_string_to_number is working with the input '$100 '.
```
>>> 
>>> convert_pay_string_to_number("$23 ") == 23000000.0
True
```
failure_message: ❌ convert_pay_string_to_number is not working for our example.
success_message: ✅ It looks like your function convert_pay_string_to_number is working with the input '$23 '.


Tests for task_10 (points: None)
```
>>> 
>>> len(some_functions)==3
True
```
failure_message: ❌ You don't seem to have 3 items in the array.
success_message: ✅ You have 3 items in the array.
```
>>> 
>>> callable(some_functions.item(0))
True
```
failure_message: ❌ The first item in your array doesn't seem to be a function.
success_message: ✅ It looks like your first item is a function.
```
>>> 
>>> # The second thing in your array may not be a function.
>>> callable(some_functions.item(1))
True
```
failure_message: ❌ The second item in your array doesn't seem to be a function.
success_message: ✅ It looks like your second item is a function.
```
>>> 
>>> # The third thing in your array may not be a function.
>>> callable(some_functions.item(2))
True
```
failure_message: ❌ The third item in your array doesn't seem to be a function.
success_message: ✅ It looks like your third item is a function.


Tests for task_11 (points: None)
```
>>> 
>>> 'Total Pay ($)' in compensation.labels
True
```
failure_message: ❌ You either didn't add the 'Total Pay ($)' column, or you mislabeled it
success_message: ✅ It looks like you created a column labeled Total Pay ($)
```
>>> 
>>> t = compensation.sort('Total Pay ($)', descending = True)
>>> t.column('Total Pay ($)').item(0) == 53250000.0
True
```
failure_message: ❌ You have the correct column in your table, but the values may be wrong. Try using the apply method with convert_pay_string_to_number function. You'll need to add that information to the table using with_column.
success_message: ✅ We checked a value in the table and it looks correct!


Tests for task_12 (points: None)
```
>>> 
>>> import math
>>> math.isclose(average_total_pay, 11445294.11764706, rel_tol = 0.1)
True
```
failure_message: ❌ It doesn't seem like you have the correct value. Try using np.mean or np.average.
success_message: ✅ Great work!


Tests for task_13 (points: None)
```
>>> 
>>> len(cash_proportion) == 102
True
```
failure_message: ❌ It doesn't seem that you have the correct number of items in the array.
success_message: ✅ You have the correct number of items in the array.
```
>>> 
>>> import math
>>> math.isclose(cash_proportion.item(0), 0.01784038, rel_tol = .001)
True
```
failure_message: ❌ It doesn't seem that you have the correct values in the array. (Don't try to fix the warning by editing the contents of the table.)
success_message: ✅ It looks like you have the correct values in the array.


Tests for task_14 (points: None)
```
>>> 
>>> "(No previous year)" not in with_previous_compensation.column("% Change")
True
```
failure_message: ❌ Make sure to remove the '(No previous year)' CEOs.
success_message: ✅ It seems that you've removed the '(No previous year)' CEOs.
```
>>> 
>>> '2014 Total Pay ($)' in set(with_previous_compensation.labels)
True
```
failure_message: ❌ You don't have the column '2014 Total Pay ($)' in with_previous_compensation.
success_message: ✅ You have the column '2014 Total Pay ($)' in with_previous_compensation.
```
>>> 
>>> import math
>>> 
>>> t = with_previous_compensation.sort("2014 Total Pay ($)", descending = True)
>>> value = t.column("2014 Total Pay ($)").item(0)
>>> math.isclose(value, 67700000.0, rel_tol = 1000)
True
```
failure_message: ❌ We checked one of your values in with_previous_compensation and it doesn't seeem correct.
success_message: ✅ We checked one of your values in with_previous_compensation and it seeems correct.
```
>>> 
>>> with_previous_compensation.num_rows==81
True
```
failure_message: ❌ You have the column, but your number of rows is off.
success_message: ✅ You have the column and the correct number of rows.


Tests for task_15 (points: None)
```
>>> 
>>> import math
>>> math.isclose(average_pay_2014, 11649176.115603436, rel_tol = 10)
True
```
failure_message: ❌ It doesn't seem that you have the correct average.
success_message: ✅ Great work!
