```python
# Initialize Otter
import otter
grader = otter.Notebook("hw03.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

<div style="text-align: center;">
    <h1>Homework 03 - Table Manipulation and Visualization</h1>
    <em>View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.</em>
</div>

**Recommended Reading**: 
* [Visualization](https://www.inferentialthinking.com/chapters/07/Visualization.html)

Please complete this notebook by filling in the cells provided. 

**Throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Moreover, please be sure to only put your written answers in the provided cells. 

Run the following cells to import the modules and do some configuration for this assignment.

```python
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
```

## Unemployment


The Federal Reserve Bank of St. Louis publishes data about jobs in the US.  Below, we've loaded data on unemployment in the United States. There are many ways of defining unemployment, and our dataset includes two notions of the unemployment rate:

1. Among people who are able to work and are looking for a full-time job, the percentage who can't find a job.  This is called the Non-Employment Index, or NEI.
2. Among people who are able to work and are looking for a full-time job, the percentage who can't find any job *or* are only working at a part-time job.  The latter group is called "Part-Time for Economic Reasons", so the acronym for this index is NEI-PTER.  (Economists are great at marketing.)

The source of the data is from [FRED](https://fred.stlouisfed.org/categories/33509).

### Task 01 📍

The data are in a CSV file called `unemployment.csv`.  Load that file into a table called `unemployment`.


_Points:_ 1

```python
unemployment = ...
unemployment
```

```python
grader.check("task_01")
```

### Task 02 📍

Sort the data in descending order by NEI, naming the sorted table `by_nei`.  Create another table called `by_nei_pter` that's sorted in descending order by NEI-PTER instead.


_Points:_ 1

```python
by_nei = ...
by_nei_pter = ...
```

```python
grader.check("task_02")
```

### Task 03 📍

Use `take` to make a table containing the data for the 10 quarters when NEI was greatest.  Call that table `greatest_nei`.

`greatest_nei` should be sorted in descending order of `NEI`. Note that each row of `unemployment` represents a quarter.


_Points:_ 2

```python
greatest_nei = ...
greatest_nei
```

```python
grader.check("task_03")
```

### Task 04 📍

It's believed that many people became PTER (recall: "Part-Time for Economic Reasons") in the "Great Recession" of 2008-2009.  NEI-PTER is the percentage of people who are unemployed (and counted in the NEI) plus the percentage of people who are PTER.  Compute an array containing the percentage of people who were PTER in each quarter.  (The first element of the array should correspond to the first row of `unemployment`, and so on.)

*Note:* Use the original `unemployment` table for this.


_Points:_ 2

```python
pter = ...
pter
```

```python
grader.check("task_04")
```

### Task 05 📍

Add `pter` as a column to `unemployment` (named "PTER") and sort the resulting table by that column in descending order.  Call the table `by_pter`.

Try to do this with a single line of code, if you can.


_Points:_ 2

```python
by_pter = ...
by_pter
```

```python
grader.check("task_05")
```

### Task 06 📍

Create a new table called `pter_over_time` that adds the `year` array and the `pter` array to the `unemployment` table. Label these columns `Year` and `PTER`.


_Points:_ 2

```python
year = 1994 + np.arange(by_pter.num_rows)/4
pter_over_time = ...
```

```python
grader.check("task_06")
```

### Task 07 📍🔎

<!-- BEGIN QUESTION -->

Use the table `pter_over_time` to create a line plot of the PTER over time using one of the table methods you've learned from the course material. Your graph should show the `Year` values on the horizontal axis and the `PTER` values on the vertical axis.

_Points:_ 2

```python
# Generate your chart in this cell
pter_over_time.plot("Year", "PTER") # SOLUTION
```

<!-- END QUESTION -->

### Task 08 📍

Were PTER rates high during the Great Recession (that is to say, were PTER rates particularly high in the years 2008 through 2011)? Assign highPTER to `True` if you think PTER rates were high in this period, and `False` if you think they weren't. 


_Points:_ 1

```python
highPTER = ...
```

```python
grader.check("task_08")
```

## Birth Rates


The following table gives census-based population estimates for each state on both July 1, 2015 and July 1, 2016. The last four columns describe the components of the estimated change in population during this time interval. **For all questions below, assume that the word "states" refers to all 52 rows including Puerto Rico & the District of Columbia.**

The data was taken from [the US Census 2010-2016 national totals data set](http://www2.census.gov/programs-surveys/popest/datasets/2010-2016/national/totals/nst-est2016-alldata.csv).

If you want to read more about the different column descriptions, click [here](http://www2.census.gov/programs-surveys/popest/datasets/2010-2015/national/totals/nst-est2015-alldata.pdf)!

The raw data is a bit messy - run the cell below to clean the table and make it easier to work with.

```python
# Don't change this cell; just run it.
pop = Table.read_table('nst-est2016-alldata.csv').where('SUMLEV', 40).select([1, 4, 12, 13, 27, 34, 62, 69])
pop = pop.relabeled('POPESTIMATE2015', '2015').relabeled('POPESTIMATE2016', '2016')
pop = pop.relabeled('BIRTHS2016', 'BIRTHS').relabeled('DEATHS2016', 'DEATHS')
pop = pop.relabeled('NETMIG2016', 'MIGRATION').relabeled('RESIDUAL2016', 'OTHER')
pop = pop.with_columns("REGION", np.array([int(region) if region != "X" else 0 for region in pop.column("REGION")]))
pop.set_format([2, 3, 4, 5, 6, 7], NumberFormatter(decimals=0)).show(5)
```

### Task 09 📍

Assign `us_birth_rate` to the total US annual birth rate during this time interval. The annual birth rate for a year-long period is the total number of births in that period as a proportion of the population size at the start of the time period.

**Hint:** Which year corresponds to the start of the time period?


_Points:_ 2

```python
us_birth_rate = ...
us_birth_rate
```

```python
grader.check("task_09")
```

### Task 10 📍

Assign `movers` to the number of states for which the **absolute value** of the **annual rate of migration** was higher than 1%. The annual rate of migration for a year-long period is the net number of migrations (in and out) as a proportion of the population size at the start of the period. The `MIGRATION` column contains estimated annual net migration counts by state.


_Points:_ 2

```python
migration_rates = ...
movers = ...
movers
```

```python
grader.check("task_10")
```

### Task 11 📍

Assign `west_births` to the total number of births that occurred in region 4 (the Western US). 

**Hint:** Make sure you double check the type of the values in the region column, and appropriately filter (i.e. the types must match!).


_Points:_ 2

```python
west_births = ...
west_births
```

```python
grader.check("task_11")
```

### Task 12 📍

Assign `less_than_west_births` to the number of states that had a total population in 2016 that was smaller than the *total number of births in region 4 (the Western US)* during this time interval.


_Points:_ 2

```python
less_than_west_births = ...
less_than_west_births
```

```python
grader.check("task_12")
```

### Task 13 📍

In the next question, you will be creating a visualization to understand the relationship between birth and death rates for the states. The annual death rate for a year-long period in a state is the total number of deaths in that period for that state as a proportion of the population size at the start of the time period for that state.

What visualization is most appropriate to see if there is an association between birth and death rates during a given time interval among the states?

1. Line Graph
<br>
2. Scatter Plot
<br>
3. Bar Chart

Assign `visualization` below to the number corresponding to the correct visualization.


_Points:_ 2

```python
visualization = ...
```

```python
grader.check("task_13")
```

### Task 14 📍🔎

<!-- BEGIN QUESTION -->

In the code cell below, create a visualization that will help us determine if there is an association between birth rate and death rate during this time interval among the states. It may be helpful to create an intermediate table here.


_Points:_ 2

```python
# Generate your chart in this cell
...
```

<!-- END QUESTION -->

### Task 15 📍

`True` or `False`: There is an association between birth rate and death rate during this time interval among the states. 

Assign `assoc` to `True` or `False` in the cell below. 


_Points:_ 1

```python
assoc = ...
```

```python
grader.check("task_15")
```

## Marginal Histograms


Consider the following scatter plot: ![](scatter.png)

The axes of the plot represent values of two variables: $x$ and $y$. 

Suppose we have a table called `t` that has two columns in it:

- `x`: a column containing the x-values of the points in the scatter plot
- `y`: a column containing the y-values of the points in the scatter plot

Below, you are given two histograms, each of which corresponds to either column `x` or column `y`. 

### Histogram A

![](var1.png)

### Histogram B

![](var2.png)

### Task 16 📍

Suppose we run `t.hist('x')`. Which histogram does this code produce? Assign `histogram_column_x` to either 1 or 2.

1. Histogram A
2. Histogram B


_Points:_ 1

```python
histogram_column_x = ...
```

```python
grader.check("task_16")
```

### Task 17 📍🔎

<!-- BEGIN QUESTION -->

State at least one reason why you chose the histogram from Task 16. Make sure to indicate which histogram you selected (ex: "I chose histogram A because ...").


_Points:_ 1

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Task 18 📍

Suppose we run `t.hist('y')`. Which histogram does this code produce? `Assign histogram_column_y` to either 1 or 2.

1. Histogram A
2. Histogram B


_Points:_ 1

```python
histogram_column_y = ...
```

```python
grader.check("task_18")
```

### Task 19 📍🔎

<!-- BEGIN QUESTION -->

State at least one reason why you chose the histogram from Task 18.  Make sure to indicate which histogram you selected (ex: "I chose histogram A because ...").


_Points:_ 1

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Submit your Homework to Canvas

Once you have finished working on the homework tasks, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the rubric to know how you will be scored for this assignment.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
3. Select the menu item "File" and "Save Notebook" in the notebook's Toolbar to save your work and create a specific checkpoint in the notebook's work history.
4. Select the menu items "File", "Download" in the notebook's Toolbar to download the notebook (.ipynb) file. 
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded .ipynb file.

**Keep in mind that the autograder does not always check for correctness. Sometimes it just checks for the format of your answer, so passing the autograder for a question does not mean you got the answer correct for that question.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 1)
```
>>> 
>>> 'unemployment' in vars() \
... and unemployment.select('Date', 'NEI', 'NEI-PTER').take(0).column(0).item(0) == '1994-01-01'
True
```
failure_message: ❌ It doesn't seem that you have imported the data correctly as a table named unemployment.
success_message: ✅ It seems that you have imported the data correctly as a table named unemployment.


Tests for task_02 (points: 1)
```
>>> 
>>> by_nei.take(0).column(0).item(0) == '2009-10-01'
True
```
failure_message: ❌ It doesn't look seem like you have the correct values in by_nei.
success_message: ✅ It seems like you have the correct values in by_nei.
```
>>> 
>>> by_nei_pter.take(0).column(0).item(0) == '2009-10-01'
True
```
failure_message: ❌ It doesn't look seem like you have the correct values in by_nei_pter.
success_message: ✅ It seems like you have the correct values in by_nei_pter.


Tests for task_03 (points: 2)
```
>>> 
>>> greatest_nei.num_rows == 10
True
```
failure_message: ❌ There are not 10 rows (quarters) in your table.
success_message: ✅ There are 10 rows (quarters)b in your table.


Tests for task_04 (points: 2)
```
>>> 
>>> import numpy as np
>>> isinstance(pter, np.ndarray)
True
```
failure_message: ❌ pter doesn't seem to be an array.
success_message: ✅ pter is an array.
```
>>> 
>>> round(pter.item(6), 4) != -1.1282
True
```
failure_message: ❌ Becareful of subtracting the arrays in the wrong order.
success_message: ✅ You didn't seem to subtract the correct arrays in the wrong order.


Tests for task_05 (points: 2)
```
>>> 
>>> by_pter.num_columns == 4
True
```
failure_message: ❌ Your table doesn't have the correct number of columns.
success_message: ✅ Your table has the correct number of columns.


Tests for task_06 (points: 2)
```
>>> 
>>> pter_over_time.take(0).column(0).item(0) == '1994-01-01'
True
```
failure_message: ❌ We checked one value in the table and it is not correct.
success_message: ✅ We checked one value in the table and it is correct.


Tests for task_08 (points: 1)
```
>>> 
>>> highPTER == True or highPTER == False
True
```
failure_message: ❌ You didn't provide a possible answer. Use True or False.
success_message: ✅ You provided a possible answer.


Tests for task_09 (points: 2)
```
>>> 
>>> 0 < us_birth_rate < 1
True
```
failure_message: ❌ Make sure you provide a value betwen 0 and 1.
success_message: ✅ You have provided a possible value.


Tests for task_10 (points: 2)
```
>>> 
>>> 0 < movers <= 52
True
```
failure_message: ❌ Make sure you provide a value betwen 0 and 52 (inclusive).
success_message: ✅ You have provided a possible value.


Tests for task_11 (points: 2)
```
>>> 
>>> 5e5 < west_births < 1e6
True
```
failure_message: ❌ We are looking for a number between 500000 and 1000000.
success_message: ✅ You have provided a possible value.


Tests for task_12 (points: 2)
```
>>> 
>>> 0 <= less_than_west_births <= 52
True
```
failure_message: ❌ Make sure you provide a value betwen 0 and 52 (inclusive).
success_message: ✅ You have provided a possibel value.


Tests for task_13 (points: 2)
```
>>> 
>>> visualization in {1, 2, 3}
True
```
failure_message: ❌ Make sure you provide 1, 2, or 3 as a value.
success_message: ✅ You have provided a possible value.


Tests for task_15 (points: 1)
```
>>> 
>>> type(assoc) is bool
True
```
failure_message: ❌ You didn't provide a possible answer. Use True or False.
success_message: ✅ You provided a possible answer.


Tests for task_16 (points: 1)
```
>>> 
>>> histogram_column_x in {1, 2}
True
```
failure_message: ❌ Make sure you assign histogram_column_x to either 1 or 2.
success_message: ✅ You provided a possible answer.


Tests for task_18 (points: 1)
```
>>> 
>>> histogram_column_y in {1, 2}
True
```
failure_message: ❌ Make sure you assign histogram_column_y to either 1 or 2.
success_message: ✅ You provided a possible answer.
