```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

<div style="text-align: center;">
    <h1>Homework 01 - Causality and Expressions</h1>
    <em>View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.</em>
</div>

## Assignment Reminders

- Make sure to run the code cell at the top of this notebook to load the autograder.
- For all problems that you must write explanations and sentences for, provide your answer in the designated space.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- We encourage you to discuss this assignment with others, but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.

## Textbook References

- [What is Data Science](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causality and Experiments](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programming in Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

## Scary Arithmetic


### Task 01 📍🔎

<!-- BEGIN QUESTION -->

An ad for ADT Security Systems says,

> "When you go on vacation, burglars go to work [...] According to FBI statistics, over 25% of home burglaries occur between Memorial Day and Labor Day."

Do the data in the ad support the claim that burglars are more likely to go to work during the time between Memorial Day and Labor Day? Consider what percentage of the year falls between these two holidays, and please explain your answer.

**Note:** You can assume that "over 25%" means only slightly over. Had it been much over, say closer to 30%, then the marketers would have said so.


_Points:_ 1

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Characters in Little Women


In lecture, we counted the number of times that the literary characters were named in each chapter of the classic book, [*Little Women*](https://www.inferentialthinking.com/chapters/01/3/1/literary-characters). In computer science, the word "character" also refers to a letter, digit, space, or punctuation mark; any single element of a text. The following code generates a scatter plot in which each dot corresponds to a chapter of *Little Women*. The horizontal position of a dot measures the number of periods in the chapter. The vertical position measures the total number of characters.

```python
# This cell contains code that hasn't yet been covered in the course,
# but you should be able to interpret the scatter plot it generates.

from datascience import *
from urllib.request import urlopen
import numpy as np
%matplotlib inline

little_women_url = 'https://www.inferentialthinking.com/data/little_women.txt'
chapters = urlopen(little_women_url).read().decode().split('CHAPTER ')[1:]
text = Table().with_column('Chapters', chapters)
Table().with_columns(
    'Periods',    np.char.count(chapters, '.'),
    'Characters', text.apply(len, 0)
    ).scatter(0)
```

### Task 02 📍

Around how many periods are there in the chapter with the most characters? Assign either 1, 2, 3, 4, or 5 to the name `characters_task_2` below.

1. 250
2. 390
3. 440
4. 32,000
5. 40,000


_Points:_ 1

```python
characters_task_2 = ...
```

```python
grader.check("task_02")
```

The test above checks that your answers are in the correct format. **This test does not check that you answered correctly**, only that you assigned a number successfully in each multiple-choice answer cell.  The same is true for all the tests in this homework, and every other homework for this class.

### Task 03 📍

Which of the following chapters has the most characters per period? Assign either 1, 2, or 3 to the name `characters_task_3` below.
1. The chapter with about 60 periods
2. The chapter with about 350 periods
3. The chapter with about 440 periods


_Points:_ 1

```python
characters_task_3 = ...
```

```python
grader.check("task_03")
```

Again, the test above checks that your answers are in the correct format, but not that you have answered correctly.

To discover more interesting facts from this plot, read [Section 1.3.2](https://inferentialthinking.com/chapters/01/3/2/Another_Kind_Of_Character.html) of the textbook.

## 3. Names and Assignment Statements


### Task 04 📍

When you run `4 = 2 + 2` in a code cell, Python produces the cryptic error message.

<img src="./task_04_error.png" width="100%" alt="The syntax error produced by Python 3.10 when running 4 = 2 + 2">

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `names_task_4` below to indicate your answer.

1. Python is smart and already knows `4 = 2 + 2`.

2. `4` is already a defined number, and it doesn't make sense to make a number be a name for something else. In Python, "`x = 2 + 2`" means "assign `x` as the name for the value of `2 + 2`."

3. It should be `2 + 2 = 4`.

4. I don't get an error message. This is a trick question.


_Points:_ 1

```python
names_task_4 = ...
```

```python
grader.check("task_04")
```

### Task 05 📍

Review the following code and the error produced when that code is executed.

<img src="./task_05_error.png" width="100%" alt="The syntax error produced by Python 3.10 when running two = 3; six = two plus two">

Choose the best explanation of what's wrong with the code and assign 1, 2, 3, or 4 to `names_task_5` below to indicate your answer.

1. The `plus` operation only applies to numbers, not the word "two".

2. The name "two" cannot be assigned to the number 3.

3. Two plus two is four, not six.

4. Python cannot interpret the name `two` followed directly by a name that has not been defined (`plus` is not defined).


_Points:_ 1

```python
names_task_5 = ...
```

```python
grader.check("task_05")
```

### Task 06 📍

Review the following code and it's associated cryptic error message.

<img src="./task_06_error.png" width="100%" alt="The syntax error produced by Python 3.10 when running x = print(5); y = x + 2">

Choose the best explanation of what's wrong with the code and assign 1, 2, or 3 to `names_task_6` below to indicate your answer.

1. Python doesn't want `y` to be assigned.

2. The `print` operation is meant for displaying values to the programmer, not for assigning values!

3. Python can’t do addition between one name and one number. It has to be 2 numbers or 2 predefined names.


_Points:_ 1

```python
names_task_6 = ...
```

```python
grader.check("task_06")
```

## 4. Job Opportunities & Education in Rural India


A [study](http://www.nber.org/papers/w16021.pdf) at UCLA investigated factors that might result in greater attention to the health and education of girls in rural India. One such factor is information about job opportunities for women. The idea is that if people know that educated women can get good jobs, they might take more care of the health and education of girls in their families, as an investment in the girls’ future potential as earners. Without the knowledge of job opportunities, the author hypothesizes that families do not invest in women’s well-being.

The study focused on 160 villages outside the capital of India, all with little access to information about call centers and similar organizations that offer job opportunities to women. In 80 of the villages chosen at random, recruiters visited the village, described the opportunities, recruited women who had some English language proficiency and experience with computers, and provided ongoing support free of charge for three years. In the other 80 villages, no recruiters visited and no other intervention was made.

At the end of the study period, the researchers recorded data about the school attendance and health of the children in the villages.

### Task 07 📍

Which statement best describes the *treatment* and *control* groups for this study? Assign either 1, 2, or 3 to the name `jobs_task_7` below.

1. The treatment group was the 80 villages visited by recruiters, and the control group was the other 80 villages with no intervention.

2. The treatment group was the 160 villages selected, and the control group was the rest of the villages outside the capital of India.

3. There is no clear notion of *treatment* and *control* group in this study.


_Points:_ 1

```python
jobs_task_7 = ...
```

```python
grader.check("task_07")
```

### Task 08 📍

Was this an observational study or a randomized controlled experiment? Assign either 1, 2, or 3 to the name `jobs_task_8` below.

1. This was an observational study.

2. This was a randomized controlled experiment.  

3. This was a randomized observational study.


_Points:_ 1

```python
jobs_task_8 = ...
```

```python
grader.check("task_08")
```

### Task 09 📍🔎

<!-- BEGIN QUESTION -->

The study reported, "Girls aged 5-15 in villages that received the recruiting services were 3 to 5 percentage points more likely to be in school and experienced an increase in Body Mass Index, reflecting greater nutrition and/or medical care. However, there was no net gain in height. For boys, there was no change in any of these measures." Why do you think the author points out the lack of change in the boys?

*Hint:* Remember the original hypothesis. The author believes that educating women in job opportunities will cause families to invest more in the women’s well-being.


_Points:_ 1

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 5. Differences between Majors


Berkeley’s Office of Planning and Analysis provides data on numerous aspects of the campus. Adapted from the OPA website, the table below displays the numbers of degree recipients in three majors in the academic years 2008-2009 and 2017-2018.

| Major                              | 2008-2009    | 2017-2018   |
|------------------------------------|--------------|-------------|
| Gender and Women's Studies         |      17      |    28       |
| Linguistics                        |      49      |    67       |
| Rhetoric                           |      113     |    56       |



### Task 10 📍

Suppose you want to find the **biggest** absolute difference between the numbers of degree recipients in the two years, among the three majors.

In the cell below, compute this value and call it `biggest_change`. Use a single expression (a single line of code) to compute the answer. Let Python perform all the arithmetic (like subtracting 49 from 67) rather than simplifying the expression yourself. The built-in `abs` function takes a numerical input and returns the absolute value. The built-in `max` function can take in 3 arguments and returns the maximum of the three numbers


_Points:_ 1

```python
biggest_change = ...
biggest_change
```

```python
grader.check("task_10")
```

### Task 11 📍

Which of the three majors had the **smallest** absolute difference? Assign `smallest_change_major` to 1, 2, or 3 where each number corresponds to the following major:

1: Gender and Women's Studies  
2: Linguistics  
3: Rhetoric

Choose the number that corresponds to the major with the smallest absolute difference.

You should be able to answer by rough mental arithmetic, without having to calculate the exact value for each major. 


_Points:_ 1

```python
smallest_change_major = ...
smallest_change_major
```

```python
grader.check("task_11")
```

### Task 12 📍

For each major, define the "relative change" to be the following: $\large{\frac{\text{absolute difference}}{\text{value in 2008-2009}} * 100}$ 

Fill in the code below such that `gws_relative_change`, `linguistics_relative_change` and `rhetoric_relative_change` are assigned to the relative changes for their respective majors.


_Points:_ 3

```python
gws_relative_change = (abs(...) / 17) * 100
linguistics_relative_change = ...
rhetoric_relative_change = ...
gws_relative_change, linguistics_relative_change, rhetoric_relative_change
```

```python
grader.check("task_12")
```

### Task 13 📍

Assign `biggest_rel_change_major` to 1, 2, or 3 where each number corresponds to to the following: 

1: Gender and Women's Studies  
2: Linguistics  
3: Rhetoric

Choose the number that corresponds to the major with the biggest relative change.


_Points:_ 1

```python
# Assign biggest_rel_change_major to the number corresponding to the major with the biggest relative change.
biggest_rel_change_major = ...
biggest_rel_change_major
```

```python
grader.check("task_13")
```

## 6. Nearsightedness Study


Myopia, or nearsightedness, results from a number of genetic and environmental factors. In 1999, Quinn et al studied the relation between myopia and ambient lighting at night (for example, from nightlights or room lights) during childhood.

### Task 14 📍🔎

<!-- BEGIN QUESTION -->

The data were gathered by the following procedure, reported in the study. "Between January and June 1998, parents of children aged 2-16 years [...] that were seen as outpatients in a university pediatric ophthalmology clinic completed a questionnaire on the child’s light exposure both at present and before the age of 2 years." Was this study observational, or was it a controlled experiment? Explain. 


_Points:_ 1

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Task 15 📍🔎

<!-- BEGIN QUESTION -->

The study found that of the children who slept with a room light on before the age of 2, 55% were myopic. Of the children who slept with a night light on before the age of 2, 34% were myopic. Of the children who slept in the dark before the age of 2, 10% were myopic. The study concluded that, "The prevalence of myopia [...] during childhood was strongly associated with ambient light exposure during sleep at night in the first two years after birth."

Do the data support this statement? You may interpret "strongly" in any reasonable qualitative way.


_Points:_ 1

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Task 16 📍🔎

<!-- BEGIN QUESTION -->

On May 13, 1999, CNN reported the results of this study under the headline, "Night light may lead to nearsightedness." Does the conclusion of the study claim that night light causes nearsightedness?


_Points:_ 1

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Task 17 📍🔎

<!-- BEGIN QUESTION -->

The final paragraph of the CNN report said that "several eye specialists" had pointed out that the study should have accounted for heredity.

Myopia is passed down from parents to children. Myopic parents are more likely to have myopic children, and may also be more likely to leave lights on habitually (since the parents have poor vision). In what way does the knowledge of this possible genetic link affect how we interpret the data from the study? 


_Points:_ 1

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 7. Studying the Survivors


### Task 18 📍

The Reverend Henry Whitehead was skeptical of John Snow’s conclusion about the Broad Street pump. After the Broad Street cholera epidemic ended, Whitehead set about trying to prove Snow wrong.  (The history of the event is detailed [here](http://www.ncbi.nlm.nih.gov/pmc/articles/PMC1034367/pdf/medhist00183-0026.pdf).)

He realized that Snow had focused his analysis almost entirely on those who had died. Whitehead, therefore, investigated the drinking habits of people in the Broad Street area who had not died in the outbreak.

What is the main reason it was important to study this group? Assign 1, 2, or 3 to `survivor_answer`.

1) If Whitehead had found that many people had drunk water from the Broad Street pump and not caught cholera, that would have been evidence against Snow's hypothesis.

2) Survivors could provide additional information about what else could have caused the cholera, potentially unearthing another cause.

3) Through considering the survivors, Whitehead could have identified a cure for cholera.


_Points:_ 1

```python
survivor_answer = ...
```

```python
grader.check("task_18")
```

Whitehead ended up finding further proof that the Broad Street pump played the central role in spreading the disease to the people who lived near it. Eventually, he became one of Snow’s greatest defenders.

## Submit your Homework to Canvas

Once you have finished working on the homework tasks, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the rubric to know how you will be scored for this assignment.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
3. Select the menu item "File" and "Save Notebook" in the notebook's Toolbar to save your work and create a specific checkpoint in the notebook's work history.
4. Select the menu items "File", "Download" in the notebook's Toolbar to download the notebook (.ipynb) file. 
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded .ipynb file.

**Keep in mind that the autograder does not always check for correctness. Sometimes it just checks for the format of your answer, so passing the autograder for a question does not mean you got the answer correct for that question.**

Tests for task_02 (points: 2)
```
>>> 
>>> characters_task_2 in {1, 2, 3, 4, 5}
True
```
failure_message: ❌ Make sure you are entering 1, 2, 3, 4, or 5 as your value.
success_message: ✅ You successfully submitted a possible value.


Tests for task_03 (points: 1)
```
>>> 
>>> characters_task_3 in {1, 2, 3}
True
```
failure_message: ❌ Make sure you are entering 1, 2, or 3 as your value.
success_message: ✅ You successfully submitted a possible value.


Tests for task_04 (points: 1)
```
>>> 
>>> names_task_4 in {1, 2, 3, 4}
True
```
failure_message: ❌ Make sure you are entering 1, 2, 3, or 4 as your value.
success_message: ✅ You successfully submitted a possible value.


Tests for task_05 (points: 1)
```
>>> names_task_5 in {1, 2, 3, 4}
True
```
failure_message: ❌ Make sure you are entering 1, 2, 3, or 4 as your value.
success_message: ✅ You successfully submitted a possible value.


Tests for task_06 (points: 1)
```
>>> 
>>> names_task_6 in {1, 2, 3}
True
```
failure_message: ❌ Make sure you are entering 1, 2, or 3 as your value.
success_message: ✅ You successfully submitted a possible value.


Tests for task_07 (points: 1)
```
>>> 
>>> jobs_task_7 in {1, 2, 3}
True
```
failure_message: ❌ Make sure you are entering 1, 2, or 3 as your value.
success_message: ✅ You successfully submitted a possible value.


Tests for task_08 (points: 1)
```
>>> 
>>> jobs_task_8 in {1, 2, 3}
True
```
failure_message: ❌ Make sure you are entering 1, 2, or 3 as your value.
success_message: ✅ You successfully submitted a possible value.


Tests for task_10 (points: 1)
```
>>> 
>>> isinstance(biggest_change,(int, float))
True
```
failure_message: ❌ Make sure you replacing the ... with an expression that evaluates to be a number.
success_message: ✅ You successfully assigned biggest_change to be number.


Tests for task_11 (points: 1)
```
>>> 
>>> isinstance(smallest_change_major, (int, float))
True
```
failure_message: ❌ Make sure you replacing the ... with an expression that evaluates to be a number.
success_message: ✅ You successfully assigned smallest_change_major to be number.


Tests for task_12 (points: 3)
```
>>> 
>>> isinstance(gws_relative_change, (int, float))
True
```
failure_message: ❌ Make sure you replacing the ... with an expression that evaluates to be a number.
success_message: ✅ You successfully assigned gws_relative_change to be number.
```
>>> 
>>> isinstance(linguistics_relative_change, (int, float))
True
```
failure_message: ❌ Make sure you replacing the ... with an expression that evaluates to be a number.
success_message: ✅ You successfully assigned linguistics_relative_change to be number.
```
>>> 
>>> isinstance(rhetoric_relative_change, (int, float))
True
```
failure_message: ❌ Make sure you replacing the ... with an expression that evaluates to be a number.
success_message: ✅ You successfully assigned rhetoric_relative_change to be number.


Tests for task_13 (points: 1)
```
>>> 
>>> biggest_rel_change_major in {1, 2, 3}
True
```
failure_message: ❌ Make sure you are entering 1, 2, or 3 as your value.
success_message: ✅ You successfully submitted a possible value.


Tests for task_18 (points: 1)
```
>>> 
>>> survivor_answer in {1, 2, 3}
True
```
failure_message: ❌ Make sure you are entering 1, 2, or 3 as your value.
success_message: ✅ You successfully submitted a possible value.
