```python
# Initialize Otter
import otter
grader = otter.Notebook("hw09.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

<div style="text-align: center;">
    <h1>Homework 9: Bootstrap, Resampling, and the CLT</h1>
    <em>View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.</em>
</div>

**Reading**: 
* [Estimation](https://www.inferentialthinking.com/chapters/13/Estimation.html)
* [Why the mean matters](https://inferentialthinking.com/chapters/14/Why_the_Mean_Matters.html)

Please complete this notebook by filling in the cells provided.  

For all problems that you must write our explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on.

Run the following code cell.

```python
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
#import warnings
#warnings.simplefilter('ignore', FutureWarning)
```

## Preliminaries

The British Royal Air Force wanted to know how many warplanes the Germans had (some number `N`, which is a *parameter*), and they needed to estimate that quantity knowing only a random sample of the planes' serial numbers (from 1 to `N`). We know that the German's warplanes are labeled consecutively from 1 to `N`, so `N` would be the total number of warplanes they have. 

We normally investigate the random variation among our estimates by simulating a sampling procedure from the population many times and computing estimates from each sample that we generate.  In real life, if the British Royal Air Force (RAF) had known what the population looked like, they would have known `N` and would not have had any reason to think about random sampling. However, they didn't know what the population looked like, so they couldn't have run the simulations that we normally do. 

Simulating a sampling procedure many times was a useful exercise in *understanding random variation* for an estimate, but it's not as useful as a tool for practical data analysis.

Let's flip that sampling idea on its head to make it practical. **Given *just* a random sample of serial numbers, we'll estimate `N`, and then we'll use simulation to find out how accurate our estimate probably is, without ever looking at the whole population.**  This is an example of *statistical inference*.

We (the RAF in World War II) want to know the number of warplanes fielded by the Germans.  That number is `N`.  The warplanes have serial numbers from 1 to `N`, so `N` is also equal to the largest serial number on any of the warplanes.

We only see a small number of serial numbers (assumed to be a random sample with replacement from among all the serial numbers), so we have to use estimation.

### Task 01 📍

Is `N` a population parameter or a statistic?  If we use our random sample to compute a number that is an estimate of `N`, is that a population parameter or a statistic?

Set `N` and `N_estimate` to either the string `"parameter"` or `"statistic"` to indicate whether each value is a parameter or a statistic.


_Points:_ 3

```python
N = ...
N_estimate = ...
```

```python
grader.check("task_01")
```

To make the situation realistic, we're going to hide the true number of warplanes from you.  You'll have access only to this random sample:

```python
observations = Table.read_table("serial_numbers.csv")
num_observations = observations.num_rows
observations
```

### Task 02 📍

The average of the sample is about half of `N`. So one way to estimate `N` is to take twice the mean of the serial numbers we see. 

1. Write a function `mean_based_estimator` that computes that statistic. It should take as its argument an array of serial numbers and return twice their mean. 
2. Use the function and the `observations` table to compute an estimate of `N` called `mean_based_estimate`.


_Points:_ 2

```python
def mean_based_estimator(nums):
    ...

mean_based_estimate = ...
mean_based_estimate
```

```python
grader.check("task_02")
```

### Task 03 📍

We can also estimate `N` by using the biggest serial number in the sample.  Compute this value and give it the name `max_estimate`.


_Points:_ 3

```python
max_estimate = ...
max_estimate
```

```python
grader.check("task_03")
```

### Task 04 📍🔎

<!-- BEGIN QUESTION -->

Let's take a look at the values of `max_estimate` and `mean_based_estimate` that we got for our dataset.

1. Which of these values is closer to the true population maximum `N`? 
1. Based off of our estimators, can we give a lower bound for what `N` must be? In other words, is there a value that `N` must be greater than or equal to?


_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

We can't just confidently proclaim that `max_estimate` or `mean_based_estimate` is equal to `N`.  What if we're really far off?  We want to get a sense of the accuracy of our estimates.

## Resampling

To do this, we'll use resampling.  That is, we won't exactly simulate the observations the RAF would have really seen.  Rather we sample from our current sample, or "resample."

Why does that make any sense?

When we try to find the value of a population parameter, we ideally would like to use the whole population. However, we often only have access to one sample and we must use that to estimate the parameter instead.

Here, we would like to use the population of serial numbers to draw more samples and run a simulation about estimates of `N`.  But we still only have our sample.  So, we **use our sample in place of the population** to run the simulation. We resample from our original sample with replacement as many times as there are elements in the original sample. This resampling technique is called *bootstrapping*. 

Note that in order for bootstrapping to work well, you must start with a large, random sample. Then the Law of Large Numbers says that with high probability, your sample is representative of the population.

### Task 05 📍

Write a function called `simulate_resample`. The function should take one argument `tbl`, which is a table like `observations`. The function should generate a resample from the observed serial numbers in `tbl`.


_Points:_ 3

```python
def simulate_resample(tbl):
    ...

simulate_resample(observations) # Don't delete this line
```

```python
grader.check("task_05")
```

We'll use many resamples at once to see what estimates typically look like.  However, we don't often pay attention to single resamples, so it's easy to misunderstand them.  Let's first answer some questions about our resample.

### Task 06 📍

Which of the following statements are true?

1. The original sample can contain serial numbers that are not in the resample.
2. Because the sample size is small, the histogram of the resample might look very different from the histogram of the original sample.
3. The resample can contain serial numbers that are not in the original sample.
4. The original sample has exactly one copy of each serial number for every German plane.
5. The resample has either zero, one, or more than one copy of each serial number.
6. The resample has exactly the same sample size as the original sample.

Assign `true_statements` to an array of the number(s) corresponding to correct statements.

*Note:* The "original sample" refers to `observations`, and the "resample" refers the output of one call of `simulate_resample()`. 


_Points:_ 2

```python
true_statements = ...
```

```python
grader.check("task_06")
```

Now, let's write a function to do many resamples at once.

### Task 07 📍

Write a function called `sample_estimates`.  It should take 3 arguments:
1. `serial_num_tbl`: A table from which the data should be sampled.  The table will look like `observations`. 
2. `statistic`: A *function* that takes in an array of serial numbers as its argument and computes a statistic from the array (i.e. returns a calculated number). 
3. `num_replications`: The number of simulations to perform.

*Hint: You should use the function `simulate_resample` which you defined in Task 05*

The function should simulate many samples **with replacement** from the given table. For each of those samples, it should compute the statistic on that sample. Then it should **return an array** containing each of those statistics.  The code below provides an example use of your function and describes how you can verify that you've written it correctly.


_Points:_ 4

```python
def sample_estimates(serial_num_tbl, statistic, num_replications):
    ...

# DON'T CHANGE THE CODE BELOW THIS COMMENT! (If you do, you will fail the hidden test)
# This is just an example to test your function.
# This should generate an empirical histogram of twice-mean-based estimates
# of N from samples of size 50 if N is 1000.  This should be a bell-shaped
# curve centered at roughly 900 with most of its mass in [800, 1200].  To verify your
# answer, make sure that's what you see!
population = Table().with_column("serial number", np.arange(1, 1000+1))
one_sample = Table.read_table("one_sample.csv") #This is a sample from the population table
example_estimates = sample_estimates(
    one_sample,
    mean_based_estimator,
    10000)
Table().with_column("mean-based estimate", example_estimates).hist(bins=np.arange(0, 1500, 25))
```

```python
grader.check("task_07")
```

Now we can go back to the sample we actually observed (the table `observations`) and estimate how much our mean-based estimate of `N` would have varied from sample to sample.

### Task 08 📍

Using the bootstrap and the sample `observations`, simulate the approximate distribution of *mean-based estimates* of `N`.  Use 7,500 replications and save the estimates in an array called `bootstrap_mean_based_estimates`.  

We have provided code that plots a histogram, allowing you to visualize the simulated estimates.


_Points:_ 2

```python
bootstrap_mean_based_estimates = ...

# Don't change the code below! This plots bootstrap_mean_based_estimates.
Table().with_column("mean-based estimate", bootstrap_mean_based_estimates).hist(bins=np.arange(0, 200, 4)) 
```

```python
grader.check("task_08")
```

### Task 09 📍

Using the bootstrap and the sample `observations`, simulate the approximate distribution of *max estimates* of `N`.  Use 7,500 replications and save the estimates in an array called `bootstrap_max_estimates`.

We have provided code that plots a histogram, allowing you to visualize the simulated estimates.


_Points:_ 2

```python
bootstrap_max_estimates = ...

# Don't change the code below! This plots bootstrap_max_estimates.
Table().with_column("max estimate", bootstrap_max_estimates).hist(bins=np.arange(0, 200, 4)) 
```

```python
grader.check("task_09")
```

### Task 10 📍🔎

<!-- BEGIN QUESTION -->

`N` was actually 150! Compare the histograms of estimates you generated in Task 08 and Task 09 and answer the following questions:

1. How does the distribution of values for the mean-based estimates differ from the max estimates? Make sure to address if both distributions contain the true max value.
2. Is it more likely to get a bad max estimate for the population max then a bad mean estimate for the population mean? Explain why you think your response is correct.


_Points:_ 3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Computing intervals

### Task 11 📍

Compute an interval that covers the middle 95% of the mean-based bootstrap estimates.  Assign your values to `left_end_1` and `right_end_1`. 

*Hint:* Use the `percentile` function! Read up on its documentation in the Python Reference.

Verify that your interval looks like it covers 95% of the area in the histogram. The red dot on the histogram is the value of the parameter (150).


_Points:_ 3

```python
left_end_1 = ...
right_end_1 = ...
print("Middle 95% of bootstrap estimates: [{:f}, {:f}]".format(left_end_1, right_end_1))

# Don't change the code below! It draws your interval and N on the histogram of mean-based estimates.
Table().with_column("mean-based estimate", bootstrap_mean_based_estimates).hist(bins=np.arange(0, 200, 4)) 
plt.plot(make_array(left_end_1, right_end_1), make_array(0, 0), color='yellow', lw=7, zorder=1)
plt.scatter(150, 0, color='red', s=30, zorder=2);
```

```python
grader.check("task_11")
```

### Task 12 📍

Write code that simulates the sampling and bootstrapping process again, as follows:

1. Generate a new set of random observations the RAF might have seen by sampling from the `population` table we have created for you below. Use the sample size `num_observations`.
2. Compute an estimate of `N` from these new observations, using `mean_based_estimator`.
3. Using only the new observations, compute 10,000 bootstrap estimates of `N`.
4. Plot these bootstrap estimates and compute an interval covering the middle 95%.

*Note:* Traditionally, when we bootstrap using a sample from the population, that sample is usually a simple random sample (i.e., sampled uniformly at random from the population without replacement). However, if the population size is big enough, the difference between sampling with replacement and without replacement is negligible. Think about why that's the case! This is why when we define `new_observations`, we sample with replacement.


_Points:_ 2

```python
population = Table().with_column("serial number", np.arange(1, 150+1))

new_observations = ...
new_mean_based_estimate = ...
new_bootstrap_estimates = ...
Table().with_column("mean-based estimate", new_bootstrap_estimates).hist(bins=np.arange(0, 252, 4))
new_left_end = ...
new_right_end = ...

# Don't change code below this line!
print("New mean-based estimate: {:f}".format(new_mean_based_estimate))
print("Middle 95% of bootstrap estimates: [{:f}, {:f}]".format(new_left_end, new_right_end))

plt.plot(make_array(new_left_end, new_right_end), make_array(0, 0), color='yellow', lw=7, zorder=1)
plt.scatter(150, 0, color='red', s=30, zorder=2);
```

```python
grader.check("task_12")
```

### Task 13 📍🔎

<!-- BEGIN QUESTION -->

1. Does the interval covering the middle 95% of the new bootstrap estimates include `N`?  
2. If you ran that cell 100 times and generated 100 intervals, how many of those intervals would you expect to include `N`?


_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Let's look at what happens when we use a small number of resamples:

<img src="smallrephist.png" width="525"/>

This histogram and confidence interval was generated using 10 resamples of `new_observations`.

### Task 14 📍🔎

<!-- BEGIN QUESTION -->

In the cell below, explain why this histogram and confidence interval look different from the ones you generated previously in Task 12 where the number of resamples was 10,000.

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## The CLT and Book Reviews

Your friend has recommended you a book, so you look for it on an online marketplace. You decide to look at reviews for the book just to be sure that it's worth buying. Let's say that on Amazon, the book only has 80% positive reviews. On GoodReads, it has 95% positive reviews. You decide to investigate a bit further by looking at the percentage of positive reviews for the book on 5 different websites that you know of, and you collect these positive review percentages in a table called `reviews.csv`.

Here, we've loaded in the table for you.

```python
reviews = Table.read_table("reviews.csv") 
reviews
```

### Task 15 📍

Calculate the average percentage of positive reviews from your sample and assign it to `initial_sample_mean`.

_Points:_ 1

```python
initial_sample_mean = ...
initial_sample_mean
```

```python
grader.check("task_15")
```

You've calculated the average percentage of positive reviews from your sample, so now you want to do some inference using this information. 

### Task 16 📍

First, simulate 5000 bootstrap resamples of the positive review percentages. For each bootstrap resample, calculate the resample mean and store the resampled means in an array called `resample_positive_percentages`. Then, plot a histogram of the resampled means.


_Points:_ 2

```python
resample_positive_percentages = make_array()

for i in np.arange(5000):
    resample = ...
    resample_avg_positive = ...
    resample_positive_percentages = ...
    
# Do NOT change these lines.
(Table().with_column("Average % of Positive Reviews in Resamples", 
                     resample_positive_percentages).hist("Average % of Positive Reviews in Resamples"))
```

```python
grader.check("task_16")
```

### Task 17 📍

What is the the shape of the empirical distribution of the average percentage of positive reviews based on our original sample? What value is the distribution centered at? Assign your answer to the variable `initial_sample_mean_distribution`--your answer should be either `1`, `2`, `3`, or `4` corresponding to the following choices:

*Hint: Look at the histogram you made in Task 2. Run the cell that generated the histogram a few times to check your intuition.*

1. The distribution is approximately normal because of the Central Limit Theorem, and it is centered at the original sample mean.
2. The distribution is not necessarily normal because the Central Limit Theorem may not apply, and it is centered at the original sample mean.
3. The distribution is approximately normal because of the Central Limit Theorem, but it is not centered at the original sample mean.
4. The distribution is not necessarily normal because the Central Limit Theorem may not apply, and it is not centered at the original sample mean.


_Points:_ 2

```python
initial_sample_mean_distribution = ...
```

```python
grader.check("task_17")
```

According to the Central Limit Theorem, the probability distribution of the sum or average of a *large random sample* drawn with replacement will be roughly normal, regardless of the distribution of the population from which the sample is drawn. 

### Task 18 📍🔎

<!-- BEGIN QUESTION -->

Note the statement about the sample being large and random. Is this sample large and random? Give a brief explanation.

*Note: The setup at the beginning of this exercise explains how the sample was gathered.*


_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Though you have an estimate of the true percentage of positive reviews (the sample mean), you want to measure how variable this estimate is. 

### Task 19 📍

Find the standard deviation of your resampled average positive review percentages, which you stored in `resample_positive_percentages`, and assign the result to the variable `resampled_means_variability`.


_Points:_ 2

```python
resampled_means_variability = ...
resampled_means_variability
```

```python
grader.check("task_19")
```

This estimate is pretty variable! To make the estimate less variable, let's say you found a way to randomly sample reputable marketplaces from across the web which sell this book. Let's say that there are up to 150 of these marketplaces. The percentages of positive reviews are loaded into the table `more_reviews`.

```python
# Just run this cell
more_reviews = Table.read_table("more_reviews.csv")
more_reviews
```

In the next few questions, we'll test an important result of the Central Limit Theorem. According to the CLT, the standard deviation of all possible sample means can be calculated using the following formula:

$$
\text{SD of all possible sample means} = \dfrac{\text{Population SD}}{\sqrt{\text{sample size}}}
$$

This formula gives us another way to approximate the SD of the sample means other than calculating it empirically. We can test how well this formula works by calculating the SD of sample means for different sample sizes.

The following code calculates the SD of sample means using the CLT and empirically for a range of sample sizes. Then, it plots a scatter plot comparing the SD of the sample means calculated with both methods. Each point corresponds to a different sample size. 

```python
# Just run this cell. It's not necessary for you to read this code, but you can do 99% of this on your own!
# Note: this cell might take a bit to run.

def empirical_sample_mean_sd(n):
    sample_means = make_array()
    for i in np.arange(500):
        sample = more_reviews.sample(n).column('Positive Review Percentage') 
        sample_mean = np.mean(sample) 
        sample_means = np.append(sample_means, sample_mean) 
    return np.std(sample_means)

def predict_sample_mean_sd(n):
    return np.std(more_reviews.column(0)) / (n**0.5)

sd_table = Table().with_column('Sample Size', np.arange(1,151))
predicted = sd_table.apply(predict_sample_mean_sd, 'Sample Size')
empirical = sd_table.apply(empirical_sample_mean_sd, 'Sample Size')
sd_table = sd_table.with_columns('Predicted SD', predicted, 'Empirical SD', empirical)
sd_table.scatter('Sample Size')
plt.ylabel("SD of Sample Mean");
```

### Task 20 📍

Assign the numbers corresponding to all true statements to an array called `sample_mean_sd_statements`. 

1. The law of large numbers tells us that the distribution of a large random sample should resemble the distribution from which it is drawn.
2. The SD of the sample means is proportional to the square root of the sample size. 
3. The SD of the sample means is proportional to 1 divided by the square root of the sample size. 
4. The law of large numbers guarantees that empirical and predicted sample mean SDs will be exactly equal to each other when the sample size is large.
5. The law of large numbers guarantees that empirical and predicted sample mean SDs will be approximately equal to each other when the sample size is large.
6. The plot above shows that as our sample size increases, our estimate for the true percentage of positive reviews becomes more accurate.
7. The plot above shows that the size of the population affects the SD of the sample means.


_Points:_ 2

```python
sample_mean_sd_statements = ...
```

```python
grader.check("task_20")
```

Often times, when conducting statistical inference, you'll want your estimate of a population parameter to have a certain accuracy. It is common to measure accuracy of an estimate using the SD of the estimate--as the SD goes down, your estimate becomes less variable. As a result, the width of the confidence interval for your estimate decreases (think about why this is true). We know from the Central Limit Theorem that when we estimate a sample mean, the SD of the sample mean decreases as the sample size increases (again, think about why this is true).  

### Task 21 📍

Imagine you are asked to estimate the true average percentage of positive reviews for this book and you have not yet taken a sample of review websites. Which of these is the best way to decide how large your sample should be to achieve a certain level of accuracy for your estimate of the true average percentage of positive reviews? Assign `sample_size_calculation` to either `1`, `2`, or `3` corresponding to the statements below.

*Note: Assume you know the population SD or can estimate it with reasonable accuracy.*
1. Take many random samples of different sizes, then calculate empirical confidence intervals using the bootstrap until you reach your desired accuracy. 
2. Use the Central Limit Theorem to calculate what sample size you need in advance. 
3. Randomly pick a sample size and hope for the best.


_Points:_ 2

```python
sample_size_calculation = ...
```

```python
grader.check("task_21")
```

## Submit your Homework to Canvas

Once you have finished working on the homework tasks, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the rubric to know how you will be scored for this assignment.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
3. Select the menu item "File" and "Save Notebook" in the notebook's Toolbar to save your work and create a specific checkpoint in the notebook's work history.
4. Select the menu items "File", "Download" in the notebook's Toolbar to download the notebook (.ipynb) file. 
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded .ipynb file.

**Keep in mind that the autograder does not always check for correctness. Sometimes it just checks for the format of your answer, so passing the autograder for a question does not mean you got the answer correct for that question.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 3)
```
>>> 
>>> N == "parameter" or N == "statistic"
True
```
failure_message: ❌ N doesn't represent a possible value. Use 'parameter' or 'statistic'.
success_message: ✅ N represents a possible value.
```
>>> 
>>> N_estimate == "parameter" or N_estimate == "statistic"
True
```
failure_message: ❌ N_estimate doesn't represent a possible value. Use 'parameter' or 'statistic'. 
success_message: ✅ N_estimate represents a possible value.


Tests for task_02 (points: 2)
```
>>> mean_based_estimator(np.array([1,2,3])) is not None
True
```
```
>>> 
>>> int(np.round(mean_based_estimator(make_array([1,2,3])))) == 4
True
```
failure_message: ❌ mean_based_estimator didn't work for our sample data.
success_message: ✅ mean_based_estimator worked for our sample data.


Tests for task_03 (points: 3)
```
>>> 
>>> type(max_estimate) in set([int, np.int32, np.int64])
True
```
failure_message: ❌ max_estimate should be an integeter.
success_message: ✅ max_estimate is an integeter.
```
>>> 
>>> max_estimate in observations.column(0)
True
```
failure_message: ❌ max_estimate should be a serial number in the observations.
success_message: ✅ max_estimate is an serial number in the observations.


Tests for task_05 (points: 3)
```
>>> 
>>> simulate_resample(observations).num_rows == 17
True
```
failure_message: ❌ Your resample should have the same number of rows as the original sample.
success_message: ✅ Your resample has the same number of rows as the original sample.
```
>>> 
>>> simulate_resample(observations).labels[0] == observations.labels[0]
True
```
failure_message: ❌ Your resample should only have the same first label.
success_message: ✅ Your resample only has the same first label.


Tests for task_06 (points: 2)
```
>>> 
>>> min(true_statements) >= 1 and max(true_statements) <= 6
True
```
failure_message: ❌ true_statements doesn't represent a possible answer.
success_message: ✅ true_statements represents a possible answer.


Tests for task_07 (points: 4)
```
>>> 
>>> len(example_estimates) == 10000
True
```
failure_message: ❌ example_estimates should contain 10,000 items.
success_message: ✅ example_estimates contains 10,000 items.
```
>>> 
>>> 850 < np.mean(example_estimates) < 1100
True
```
failure_message: ❌ The mean of example_estimates should be between 850 and 1100.
success_message: ✅ The mean of example_estimates is between 850 and 1100.
```
>>> 
>>> np.count_nonzero(np.diff(example_estimates)) >= 1
True
```
failure_message: ❌ example_estimates contains the same values.
success_message: ✅ example_estimates contains different values.


Tests for task_08 (points: 2)
```
>>> 
>>> 118 < np.mean(bootstrap_mean_based_estimates) < 126
True
```
failure_message: ❌ The mean of bootstrap_mean_based_estimates should be between 118 and 126.
success_message: ✅ The mean of bootstrap_mean_based_estimates is between 118 and 126.


Tests for task_09 (points: 2)
```
>>> 
>>> max(bootstrap_max_estimates) <= 135
True
```
failure_message: ❌ The max value in bootstrap_max_estimates should be at most 135.
success_message: ✅ The max value in bootstrap_max_estimates is at most 135.


Tests for task_11 (points: 3)
```
>>> 
>>> type(left_end_1) in set([float, np.float32, np.float64]) and type(right_end_1) in set([float, np.float32, np.float64])
True
```
failure_message: ❌ left_end_1 and right_end_1 should be floats.
success_message: ✅ left_end_1 and right_end_1 are floats.


Tests for task_12 (points: 2)
```
>>> 
>>> type(new_mean_based_estimate) in set([float, np.float32, np.float64])
True
```
failure_message: ❌ new_mean_based_estimate should be a float.
success_message: ✅ new_mean_based_estimate is a float.
```
>>> 
>>> type(new_left_end) in set([float, np.float32, np.float64]) and type(new_right_end) in set([float, np.float32, np.float64])
True
```
failure_message: ❌ new_left_end and new_right_end should be floats.
success_message: ✅ new_left_end and new_right_end are floats.


Tests for task_15 (points: 1)
```
>>> 
>>> np.round(initial_sample_mean) == 74
True
```
failure_message: ❌ initial_sample_mean should be about 74
success_message: ✅ initial_sample_mean is about 74.


Tests for task_16 (points: 2)
```
>>> 
>>> len(resample_positive_percentages) == 5000
True
```
failure_message: ❌ resample_positive_percentages should contain 5,000 items.
success_message: ✅ resample_positive_percentages contains 5,000 items.


Tests for task_17 (points: 2)
```
>>> 
>>> initial_sample_mean_distribution in {1, 2, 3, 4}
True
```
failure_message: ❌ initial_sample_mean_distribution should name 1, 2, 3, or 4.
success_message: ✅ initial_sample_mean_distribution is a possible value.


Tests for task_19 (points: 2)
```
>>> 
>>> type(resampled_means_variability) in set([float, np.float32, np.float64])
True
```
failure_message: ❌ resampled_means_variability should be a float.
success_message: ✅ resampled_means_variability is a float.


Tests for task_20 (points: 2)
```
>>> 
>>> type(sample_mean_sd_statements) == np.ndarray
True
```
failure_message: ❌ sample_mean_sd_statements should be an array.
success_message: ✅ sample_mean_sd_statements is an array.


Tests for task_21 (points: 2)
```
>>> 
>>> type(sample_size_calculation) == int
True
```
failure_message: ❌ sample_size_calculation should be an integer.
success_message: ✅ sample_size_calculation is an integer.
