```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

<div style="text-align: center;">
    <h1>Homework 5: Applying Functions and Iteration</h1>
    <em>View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.</em>
</div>

Please complete this notebook by filling in the cells provided. 

**Throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Moreover, please be sure to only put your written answers in the provided cells.**

Run the following cell to import the modules and settings needed for this assignment.

```python
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## 2022-23 CCSF Football Season

You are going to analyze how well the CCSF football team performed in the 2022-23 season. A football game is divided into four periods, called quarters. The number of points CCSF scored in each quarter, and the number of points their opponent scored in each quarter are stored in a table called `ccsf_fb.csv`.

Note that the 2022-23 season data was collected from [the CCSF Athletics website](https://ccsfathletics.com/sports/fball/2022-23/teams/sanfrancisco).

```python
games = Table().read_table("ccsf_fb.csv")
games.show()
```

Let's start by finding the total points each team scored in a game.

### Task 01 📍

Write a function called `sum_scores`.  It should take five arguments, where each argument is the team's quarter or overt time score. It should return the team's total score for that game.

_Points:_ 2

```python
def sum_scores(...):
     """Returns the total score calculated by adding up the score of each quarter and the over time"""
     ...

sum_scores(14, 7, 3, 0, 3) #DO NOT CHANGE THIS LINE
```

```python
grader.check("task_01")
```

### Task 02 📍

Create a new table `final_scores` with three columns in this *specific* order: `Opponent`, `CCSF Score`, `Opponent Score`. You will have to create the `CCSF Score` and `Opponent Score` columns. Use the function `sum_scores` you just defined in the previous question for this problem.

*Hint:* If you want to apply a function that takes in multiple arguments, you can pass multiple column names as arguments in `tbl.apply()`. The column values will be passed into the corresponding arguments of the function. Take a look at the python reference for syntax.

*Tip:* If you’re running into issues creating final_scores, check that `ccsf_scores` and `opponent_scores` output what you want. 



_Points:_ 3

```python
ccsf_scores = ...
opponent_scores = ...
final_scores = ...
final_scores
```

```python
grader.check("task_02")
```

We can get specific row objects from a table. You can use `tbl.row(n)` to get the `n`th row of a table. `row.item("column_name")` will allow you to select the element that corresponds to `column_name` in a particular row. Here's an example:

```python
# Just run this cell
games.row(10)
```

```python
# Just run this cell
games.row(10).item("CCSF 4Q")
```

### Task 03 📍

We want to see for a particular game whether or not CCSF won. Write a function called `did_ccsf_win`.  It should take one argument: a row object from the `final_scores` table. It should return either `True` if CCSF's score was greater than the Opponent's score, and `False` otherwise.



_Points:_ 3

```python
def ...(row):
    ...
```

```python
grader.check("task_03")
```

### Task 04 📍

Unfortunately, CCSF did not win against every opponent during the 2022-23 season. Using the `final_scores` table, assign `results` to an array of `True` and `False` values that correspond to whether or not CCSF won. Add the `results` array to the `final_scores` table, and assign this to `final_scores_with_results`. Then, respectively assign the number of wins and losses CCSF had to `ccsf_wins` and `ccsf_losses`.

*Hint*: When you only pass a function name and no column labels through `tbl.apply()`, the function gets applied to every row in `tbl`



_Points:_ 4

```python
results = ...
final_scores_with_results = ...
ccsf_wins = ...
ccsf_losses = ...

# Don't delete or edit the following line:
print(f"In the 2022-23 Season, CCSF Football won {ccsf_wins} games and lost {ccsf_losses} games. Go RAMS!")
```

```python
grader.check("task_04")
```

### Task 05 📍

Sometimes in football the two teams are equally matched and the game is quite close. Other times, it is a blowout, where the winning team wins by a large margin of victory. Let's define a **big win** to be a game in which the winning team won by more than 10 points. Use your `final_scores` table to assign `big_wins` to an array of team names that CCSF had big wins against during the 2022-23 football season. You may find the `is_big_win` function defined below helpful to you.


_Points:_ 5

```python
def is_big_win(row):
    '''Return a boolean to describe whether or not a game (row) is a big win'''
    score_diff = row.item("CCSF Score") - row.item("Opponent Score")
    
    if score_diff > 10:
        return True
    else:
        return False

big_wins = ...

for row in ...: # Iterate through rows of final_scores table
    opponent =  ...
    if ...: # You should use the above function
        big_wins = np.append(big_wins, ...)

big_wins
```

```python
grader.check("task_05")
```

## Unrolling `for` Loops

### Task 06 📍🔎

<!-- BEGIN QUESTION -->

"Unrolling" a `for` loop means to manually write out all the code that it executes.  The result is code that does the same thing as the loop, but without the structure of the loop.  For example, for the following loop:

    for num in np.arange(3):
        print("The number is", num)
        
The unrolled version would look like this:

    print("The number is", 0)
    print("The number is", 1)
    print("The number is", 2)

Unrolling a `for` loop is a great way to understand what the loop is doing during each step. In this exercise, you'll practice unrolling a `for` loop.

In the question below, write code that does the same thing as the following, but with the `for` loop unrolled. Be careful of the spacing on this code segment!

```python
for joke_iteration in np.arange(3):
    print("Knock, knock.")
    print("Who's there?")
    print("Banana.")
    print("Banana who?")
print("Knock, knock.")
print("Who's there?")
print("Orange.")
print("Orange who?")
print("Orange you glad I didn't say banana?")
```

_Points:_ 1

```python
...
```

<!-- END QUESTION -->

## Submit your Homework to Canvas

Once you have finished working on the homework tasks, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the rubric to know how you will be scored for this assignment.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
3. Select the menu item "File" and "Save Notebook" in the notebook's Toolbar to save your work and create a specific checkpoint in the notebook's work history.
4. Select the menu items "File", "Download" in the notebook's Toolbar to download the notebook (.ipynb) file. 
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded .ipynb file.

**Keep in mind that the autograder does not always check for correctness. Sometimes it just checks for the format of your answer, so passing the autograder for a question does not mean you got the answer correct for that question.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 2)
```
>>> 
>>> sum_scores(14, 7, 3, 0, 3) == 27 and sum_scores(2, 3, 6, 1, 0) == 12
True
```
failure_message: ❌ Your function doesn't seem to work correctly with our 4 quarter scores and over time score.
success_message: ✅ Your function works with our 4 quarter scores and over time score.


Tests for task_02 (points: 3)
```
>>> 
>>> final_scores.num_columns == 3
True
```
failure_message: ❌ final_scores does not have the correct number of columns.
success_message: ✅ final_scores has the correct number of columns.
```
>>> 
>>> set(['Opponent', 'CCSF Score', 'Opponent Score']) == set(final_scores.labels)
True
```
failure_message: ❌ final_scores does not have the correct labels.
success_message: ✅ final_scores has the correct labels.


Tests for task_03 (points: 3)
```
>>> 
>>> 'did_ccsf_win' in globals() and callable(did_ccsf_win)
True
```
failure_message: ❌ You didn't create a function called did_ccsf_win.
success_message: ✅ You've created a function called did_ccsf_win.
```
>>> 
>>> did_ccsf_win(final_scores.row(1)) == False
True
```
failure_message: ❌ Your function doesn't that CCSF won against Sacramento City.
success_message: ✅ Your function shows that CCSF won against Sacramento City.


Tests for task_04 (points: 4)
```
>>> 
>>> 0 <= ccsf_wins <= 11
True
```
failure_message: ❌ There were only 13 games, so CCSF won 0, 1, 2, ..., or 11 games.
success_message: ✅ ccsf_wins is a possible value
```
>>> 
>>> 0 <= ccsf_losses <= 11
True
```
failure_message: ❌ There were only 13 games, so CCSF lost 0, 1, 2, ..., or 11 games.
success_message: ✅ ccsf_losses is a possible value


Tests for task_05 (points: 5)
```
>>> 
>>> row.item('CCSF Score') == 34
True
```
failure_message: ❌ One of the row items we checked doesn't the correct CCSF Score.
success_message: ✅ One of the row items we checked has the correct CCSF Score.
```
>>> 
>>> 0 <= len(big_wins) <= 11
True
```
failure_message: ❌ There is not a possible value for the number items in big_wins.
success_message: ✅ There is a possible value for the number items in big_wins.
```
>>> 
>>> type(big_wins.item(0)) == str
True
```
failure_message: ❌ The first item in big_wins is not a string.
success_message: ✅ The first item in big_wins is a string.
